Single-cell transcriptomics identifies an increased Ly6G- neutrophil population and accentuated T-cell cytotoxicity in tobacco flavored e-cigarette exposed mouse lungs

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

Summary:

The authors tackled the public concern about E-cigarettes among young adults by examining the lung immune environment in mice using single-cell RNA sequencing, discovering a subset of Ly6G- neutrophils with reduced IL-1 activity and increased CD8 T cells following exposure to tobaccoflavored e-cigarettes. Preliminary serum cotinine (nicotine metabolite) measurements validated the effective exposure to fruit, menthol, and tobacco-flavored e-cigarettes with air and PG:VG serving as control groups. They also highlighted the significance of metal leaching, which fluctuated over different exposure durations to flavored e-cigarettes, underscoring the inherent risks posed by these products. The scRNAseq analysis of e-cig exposure to flavors and tobacco demonstrated the most notable differences in the myeloid and lymphoid immune cell populations. Differentially expressed genes (DEGs) were identified for each group and compared against the air control. Further subclustering revealed a flavor-specific rise in Ly6G- neutrophils and heightened activation of cytotoxic T cells in response to tobacco-flavored e-cigarettes. These effects varied by sex, indicating that immune changes linked to e-cig use are dependent on gender. By analyzing the expression of various genes and employing gene ontology and gene enrichment analysis, they identified key pathways involved in this immune dysregulation resulting from flavor exposure. Overall, this study affirmed that e-cigarette exposure can suppress the neutrophil-mediated immune response, subsequently enhancing T cell toxicity in the lung tissue of mice.

Strengths:

This study used single-cell RNA sequencing to comprehensively analyze the impact of e-cigarettes on the lung. The study pinpointed alterations in immune cell populations and identified differentially expressed genes and pathways that are disrupted following e-cigarette exposure. The manuscript is well written, the hypothesis is clear, the experiments are logically designed with proper control groups, and the data is thoroughly analyzed and presented in an easily interpretable manner. Overall, this study suggested novel mechanisms by which e-cigs impact lung immunity and created a dataset that could benefit the lung immunity field.

Weaknesses:

The authors included a valuable control group - the PG:VG group, since PG:VG is the foundation of the e-liquid formulation. However, most of the comparative analyses use the air group as the control. Further analysis comparing the air group to the PG:VG group, and the PG:VG group to the individual flavored e-cig groups will provide more clear insights into the true source of irritation. This is done for a few analyses but not consistently throughout the paper. Flavor-specific effects should be discussed in greater detail. For example, Figure 1E shows that the Fruit flavor group exhibits more severe histological pathology, but similar effects were not corroborated by the singlecell data.

We thank the reviewer for this query. We agree that PG:VG group is the foundation of the e-liquid formulation and hence comparisons with this group are of significance to understand the effect of individual flavors on the cell population. Though we compared the flavored e-cig groups with PG:VG group, we did not discuss it in detail within the manuscript to avoid confusions in interpretation for this study. However, we have now included the comparisons with the PG:VG group as a Supplement File S13-S18 in our revised manuscript to facilitate proper interpretation of our omics data to interested readers.

While we agree that flavor-specific effects might be of interest, we did not delve into exploring them in detail as the fruit flavor e-liquids have now been regulated/banned from sale in the US. Thus, from regulatory point of view, the effects of tobacco-flavored e-liquids hold most interest. Since at the time of conducting this study, fruit flavors were in the market, we have still included the data. However, studying it further was not the focus of this work.

The characterization of Ly6g+ vs Ly6g- neutrophils is interesting and potentially very impactful. Key results like this from scRNAseq analyses should be validated by qPCR and flow cytometry.

Also, a recent study by Ruscitti et al reported Ly6g+ macrophages in the lung which can potentially confound the cell type analysis. A more detailed marker gene and sub-population analysis of the myeloid clusters could rule out this potential confounding factor.

We agree with the reviewer that the loss of Ly6G on neutrophils is a very interesting finding and we have designed a neutrophil specific experiment to study the impact of e-cig exposure on neutrophil maturation and function which will be discussed in subsequent work by our group. To address the concerns raised by the reviewer, we stained the lung tissue samples from air-and tobacco flavored e-cig aerosol exposed mouse lungs with Ly6G and S100A8 (universal marker for neutrophil) to see the infiltration of Ly6G+ vs Ly6G- neutrophils within the lungs of exposed and unexposed mice. Results from this study showed that exposure to tobacco-flavored e-cig aerosol affects the neutrophil population within the mouse lungs. In fact, the changes were more pronounced for female mice. The data have now been shown in Figure 4.

Reviewer #2 (Public review):

This study provides some interesting observations on how different flavors of e-cigarettes can affect lung immunology, however there are numerous flaws including a low number of replicates and a lack of effective validation methods which reduces the robustness and rigor of the findings.

Strengths:

The strength of the study is the successful scRNA-seq experiment which gives good preliminary data that can be used to create new hypotheses in this area.

Weaknesses:

The major weakness is the low number of replicates and the limited analysis methods. Two biological n per group is not acceptable to base any solid conclusions. Any validatory data was too little (only cell % data) and did not always support the findings (e.g. Figure 4D does not match 4C). Often n seems to be combined and only one data point is shown, it is not at all clear how the groups were analyzed and how many cells in each group were compared.

We thank the reviewer for recognizing the strengths of this manuscript while pointing out the errors to allow us to improve our analyses. We understand that the low number of replicates in this work makes the analyses difficult to draw solid conclusions, but this was a pilot study to identify the changes in the mouse lung upon acute exposures to flavored e-cig aerosols at a single cell level. So far, the e-cig field has been primarily focused on conducting toxicological studies to help regulatory bodies to set standards and enforce laws to better regulate the manufacture, sale and distribution of e-cig products. However, adolescents and young adults are still getting access to these products, and there is little to no understanding of how this may affect the lung health upon acute and chronic exposures. Single cell technology is a powerful tool to analyze the gene expression changes within cell populations to study cell heterogeneity and function. Yet, it is a costly tool owing to which conducting such analyses on large sample sizes is not ideal. This pilot study was designed to get some initial leads for our future studies involving larger sample sizes and chronic exposures. However, due to the vast information that is provided by a single cell RNA sequencing experiment, we intend to share it with a larger audience to support research and further study in this area. We understand that the validations are limited in our current work and so we have now conducted coimmunostaining to validate the Ly6G+ and Ly6G- neutrophil population. We have now included single cell findings with the validating experiments using classical methods of experimentation including ELISA, immunostaining or flow cytometry and revamped the whole manuscript. However, it is important to mention that such validations are sometimes challenging as many of these techniques still investigate the tissue while the changes shown in single cell analyses are mainly pertaining to a single cell type. This could be well-understood by looking at the flow cytometry results for neutrophils where we use Ly6G as a marker to stain for neutrophils which is only found in mature neutrophil population.

Only 71,725 cells mean only 7,172 per group, which is 3,586 per animal - how many of these were neutrophils, T-cells, and macrophages? This was not shown and could be too low.

We do agree that the number of cells could be too low. To avoid this, we did not study gene expression variations at the finest level of cell identity. We classified the cell clusters into general annotations -myeloid, lymphoid, endothelial, stromal and epithelial- and identified the changes in the gene expressions. Of these, only two clusters (myeloid and lymphoid) with more than ~1000 cells per cell type per group were studied in detail. We have included the cell count information to allow better interpretation of our results in the revised manuscript. For a single cell point of view, a cell count of ~3500 each with over 20000 features (genes) has good statistical strength and merit in our opinion.

The dynamic range of RNA measurement using scRNA seq is known to be limited - how do we know whether genes are not expressed or just didn't hit detection? This links into the Ly6G negative neutrophil comment, but in general, the lack of gene expression in this kind of data should be viewed with caution, especially with a low n number and few cells.

This is a well-taken point, and we thank the reviewer for this comment. We agree that the dynamic range RNA measurement is limited low cell numbers that could lead to bias. However, none of the clusters with counts lower than 150 were included for differential gene analyses. To avoid confusion, we now show immunofluorescence results to validate the findings. We are certain that with the inclusion of these validation experiments, will convince the reviewer about the loss of Ly6G marker from neutrophils and lack of proper neutrophilic response in exposed mouse lungs as compared to the controls.

There is no rigorous quantification of Ly6G+ and Ly6G- cells int he flow cytometry data.

We understand that flow-based quantification of our scRNA seq findings would be interesting. However, flow cytometry and single cell suspension to perform sequencing were performed parallelly for this study. We used a basic flow panel using single markers to identify individual immune cell type. We did identify changes in the Ly6G population in our treated and control samples using scRNA seq and intend to exclude it as a marker for our future studies using flow cytometry. Unfortunately, the same analyses could not be performed for the current batch of samples. We have now included results from IHC staining to identify the Ly6G+ and Ly6G- population in the lung tissues from control and treated mice in revised manuscript to address some of the concerns raised here.

Eosinophils are heavily involved in lung biology but are missing from the analysis.

We use RBC lysis buffer to remove the excess RBCs during lung digestion for preparation of single cell suspension for scRNA seq in this study. Reports suggest that RBC lysis could adversely affect the eosinophil number and function. We did not identify any cell cluster, representing markers for eosinophils through our scRNA seq data and we believe that our lung digestion protocol could be the reason for it. We have studied the eosinophil changes through flow cytometry in these samples and have found significant changes as well. However, due to our inability to find cell clusters for eosinophil through scRNA seq data, we did not include these results in the final manuscript previously. To avoid confusion and maintain transparency, we have now included the changes in eosinophils through flow cytometry in revised manuscript (Figure S4).

The figures had no titles so were difficult to navigate.

We have now revamped the figures to make it easier for the readers to navigate.

PGVG is not defined and not introduced early enough.

We have made the necessary changes in the revised manuscript.

Neutrophils are not well known to proliferate, so any claims about proliferation need to be accompanied by validation such as BrdU or other proliferation assays.

We have now removed the cell cycle scoring information from the revised manuscript. Performing BrDU assay was not possible for these tissues due to limited samples and resources. However, we may consider performing it in our future studies.

It was not clear how statistics were chosen and why Table S2 had a good comparison (two-way ANOVA with gender as a variable) but this was not used for other data particularly when looking at more functional RNA markers (Table S2 also lacks the interaction statistic which is most useful here).

We have now included the two-way ANOVA statistics (Supplementary File S3) for other data included in the revised manuscript. It is important to note that since we did not identify any significant changes upon two-way ANOVA, the interaction statistics were not available for the abovementioned statistical test. We have included the interaction information wherever available.

Many statistics are only vs air control, but it would be more useful as a flavor comparison to see these vs PGVG. In some cases, the carrier PGVG looks worse than some of the flavors (which have nicotine).

While we agree with this comment of the reviewer, comparisons with PG:VG were not included due to the low cell numbers for PG:VG samples obtained following quality control and filtering of scRNA seq analyses. However, considering the reviewer’s question we still include the details of comparisons with PG:VG included as supplementary files S13-S18 in the revised manuscript.

The n number is a large issue, but in Figures such as 4, 6, and 7 it could be a bigger factor. The number of significant genes identified has been determined by chance rather than any real difference, e.g. Is Il1b not identified in Fruit flavor vs air because there wasn't enough n, while in Air vs Tobacco, it randomly hit the significance mark. This is but an example of the problems with the analysis and conclusions.

While we agree in part with the concern raised here. In our opinion, an omics study is not necessarily aimed at finding the changes at transcript level with absolute certainty, but rather to identify probable cell and gene targets to validate with subsequent work. We did not claim that our findings are absolute outcomes but rather add the limitation of sample number and need for further research at every step. The strength of this work is to be the first study of its kind looking at changes in the lung cell population at single cell level upon e-cig aerosol exposure. This study has provided us with interesting gene and cell targets that we are now validating with future work. We still strongly believe that a dataset like this is a useful resource for a wider audience.

The data in Figure 7A is confusing, if this is a comparison to air, then why does air vs air not equal 1? Even if this was the comparison to the average of air between males and females, then this doesn't explain why CCL12 is >1 in both. Is this z-score instead? Regardless the data is difficult to interpret in this format.

We have now changed the format of data representation in the figure.

Individual n was not shown for almost all experiments - e.g. Figure 1D - what is this representative of? Figure 2D - is this bulk-grouped data for all cells and all mice? The heatmaps are also pooled from 2n and don't show the variability.

Wherever needed, the n number has been included in the Figure legend. Additionally, the n number is shown in Figure 1A. However, with respect to the second comment we would like to differ from the reviewer’s opinion. Each scRNA seq data had 2 samples – one for male and another for female which has been clearly shown in the current figures. The pooling of cells as mentioned in the comment happened at the stage of preparation of cell suspension from each sex/group at the start of the sequencing. We show the results of the pooled sample showing the variability amongst pooled samples, which we acknowledge is a shortcoming of our work. In terms of representation of the heat maps and data analyses we have included all the needed information to uphold transparency of our study design and data visualization for each figure and would like to stick to the current representations. However, validation cohort does not involve any pooling of sample and still agrees with most of the deductions made from this study. So we are confident that no over statements have been made in this work and we still provide a useful dataset to inform future research in this area.

Reviewer #3 (Public review):

This work aims to establish cell-type specific changes in gene expression upon exposure to different flavors of commercial e-cigarette aerosols compared to control or vehicle. Kaur et al. conclude that immune cells are most affected, with the greatest dysregulation found in myeloid cells exposed to tobacco-flavored e-cigs and lymphoid cells exposed to fruit-flavored e-cigs. The up-and-downregulated genes are heavily associated with innate immune response. The authors suggest that a Ly6G-deficient subset of neutrophils is found to be increased in abundance for the treatment groups, while gene expression remains consistent, which could indicate impaired function. Increased expression of CD4+ and CD8+ T cells along with their associated markers for proliferation and cytotoxicity is thought to be a result of activation following this decline in neutrophil-mediated immune response.

Strengths:

(1) Single-cell sequencing data can be very valuable in identifying potential health risks and clinical pathologies of lung conditions associated with e-cigarettes considering they are still relatively new.

(2) Not many studies have been performed on cell-type specific differential gene expression following exposure to e-cig aerosols.

(3) The assays performed address several factors of e-cig exposure such as metal concentration in the liquid and condensate, coil composition, cotinine/nicotine levels in serum and the product itself, cell types affected, which genes are up- or down-regulated and what pathways they control.

(4)Considerations were made to ensure clinical relevance such as selecting mice whose ages corresponded with human adolescents so that the data collected was relevant.

Weaknesses:

The exposure period of 1 hour a day for 5 days is not representative of chronic use and this time point may be too short to see a full response in all cell types. The experimental design is not well-supported based on the literature available for similar mouse models.

This study was not designed to study the effects of chronic exposures on lung tissues. We were interested in delineating the effect of acute exposures for which the proposed study design was chosen. Previous work by our group has performed similar exposures and has been well received by the community. We understand that chronic exposures will be interesting to look at, but that was beyond the scope of this pilot study. Longer / chronic exposures will be conducted considering disease modifying effects of e-cigarettes.

Several claims lack supporting evidence or use data that is not statistically significant. In particular, there were no statistical analyses to compare results across sex, so conclusions stating there is a sex bias for things like Ly6G+ neutrophil percentage by condition are observational.

We thank the reviewer for this observation, and we have now included the necessary validations and details of the sex-based statistical analyses in the revised version of this manuscript.

Statistical analyses lack rigor and are not always displayed with the most appropriate graphical representation.

We thank the reviewer and have included all the necessary statistical details with more details in the revised manuscript.

Overall, the paper and its discussion are relatively limited and do not delve into the significance of the findings or how they fit into the bigger picture of the field.

As pointed out by the reviewers themselves the strength of this work is in the first ever scRNA seq analyses of mice exposed to differently flavored e-cig aerosols in vivo. We also show cellspecific differential gene expressions and address some of the major queries made around e-cig research including release of metals on a day-to-day basis from the same coil. The limited sample number makes it difficult to draw solid conclusions from this work, which has been discussed as a shortcoming. Nevertheless, the major strength of this work is not in identifying specific trends, but rather to determine the possible cell and gene targets to expand the study for longer (chronic) exposures with a larger sample group. We have mentioned the significance of the study with respect to vaping effects on cellular heterogeneity leading to deleterious effects.

The manuscript lacks validation of findings in tissue by other methods such as staining.

We have now included some validation experiments and revamped the revised manuscript to support scRNA seq findings.

This paper provides a foundation for follow-up experiments that take a closer look at the effects of e-cig exposure on innate immunity. There is still room to elaborate on the differential gene expression within and between various cell types.

We thank the reviewer for this observation. The cell numbers for some cell clusters (especially epithelial cells) were too low. So, though we have performed the differential gene expression analyses on all the cell clusters, we refrained from discussing it in the manuscript to avoid over interpretation of our results. Only clusters with high enough (> 150) cells per sex per group were used to plot the heatmaps. We have now included the cell numbers for each cell type in the revisions to allow better interpretation of our data. Furthermore, the raw data from this study will be freely available to the public upon publication of this manuscript. This would enable the interested readers to access the raw data and study the cell types of interest in detail based on their study requirements. This data will be a useful resource for all in this community to inform and design future studies.

Recommendation For The Author:

Major comments

Mouse experiments are extremely variable and an n of 2 is not enough. Because of the complexity of separating male and female mice, the analyses are not adequately powered to support conclusions. The two-way ANOVA style approach to consider sex as a separate variable was a great idea in Table S2 - but this was not used elsewhere, and there is a need to show the interaction statistic (which would say if there is a flavor effect dependent on sex).

We thank the reviewers for this recommendation. We agree that the experiments are highly variable. However, it is not merely an outcome of a small sample size (which we address as one of the limitations). What is important to mention here is the fact that validating results from single cell technologies using regular molecular biology techniques is challenging and may not completely align. It is because we are comparing single cell population in the former and a heterogeneous cell population in latter. However, considering this comment, we have now toned down our conclusions and performed some extra experiments to validate single cell findings. We also provide the results from two-way ANOVA statistics for all the figures/experiments performed in this work.

More validatory data with PCR, immunostaining, and flow cytometry would be very helpful. This includes validating the neutrophil functional and phenotype data and the T-cell data by flow cytometry.

To validate the presence of Ly6G+ and Ly6G- neutrophil population, we performed coimmunostaining experiments and proved that exposure to tobacco-flavored e-cig aerosols results in increase in cell percentages of two neutrophil population in female mice. We also re-analyzed our Flow cytometry data to align with scRNA seq results. Multiplex protein assay was another technique used to show altered innate/adaptive immune responses upon exposure to differently flavored e-cig aerosol. Of note, considering the short duration of exposure we did not identify significant changes in cell numbers or inflammatory responses. But we have now validated our scRNA seq results using various techniques to draw meaningful conclusions.

The in vivo experimental design seems to model very short-term exposure. In the literature, including the papers cited in the references, much longer time points are used, extending from several weeks to months of exposure. There seem to be few examples of papers using 5-day exposure and those that do are inspired by traditional cigarette smoke rather than e-cig aerosols or model acute exposure by making the daily duration longer. It is important to consider the possibility that the greatest number of up- or down-regulated genes are found in immune cell populations solely because they are the first to be affected by e-cig exposure and the other cell types just do not have time to become dysregulated in 5 days.

We thank the reviewers for this comment. We do not refute the fact that our observations of major changes in the immune cell population are due to the short duration of exposure. This was one of the first studies using single cell technologies to look at cell specific changes in the mouse lungs exposed to e-cig aerosols. However, the future experiments being conducted in our lab are using more controlled approach to mimic chronic exposures to e-cig aerosols to identify changes in other cell types and long-term effects of e-cig exposures in vivo. However, since this was not the focus of this work, we have not discussed it in detail.

The validity of the claims pertaining to septal thickening and mean linear intercept (MLI) are questionable due to the poor lung inflation of the treatment group, which the authors acknowledge. Thus, MLI cannot be accurately used. It is contradictory to state that the fruit-flavored treatment group presented challenges with inflation but then concluded that there is a phenotype. In addition, inflation with low-melting agarose is not an ideal method because it does not use a liquid column to maintain constant pressure. For these metrics to be used and evaluated, it is imperative that all lobes are properly inflated. Therefore, these data should either be repeated or removed.

We agree with this critique and have removed the MLI quantification from the revised manuscripts, we also do not make claims regarding much histological changes upon exposure. We suggest further work in future to get better understanding of the effect of differently flavored e-cig aerosol exposure on mouse lungs.

What is the purpose of analyzing cell cycle scores? Why is it relevant that neutrophils are in G2M-phase? Figure 3B shows that neutrophils are clearly in both G1- and G2M-phase and this cluster includes both Ly6G+ and Ly6G- subsets, so it does not seem accurate to claim that they are in the G2M-phase of the cell cycle, nor does it reveal anything novel about Ly6G- neutrophils. Is it possible that the cell cycle score is noting a point in differentiation when neutrophils acquire/begin expressing Ly6G? Ly6G expression in neutrophils has been found to be associated with differentiation and maturation. To rule out the possibility that this is a cell state being identified, differential gene expression between the 2 neutrophil subsets should be shown in a volcano plot. It would also be useful to stain for Ly6G+/- neutrophils using either IF or RNAscope to prove they are present. If the claim is that Ly6G- neutrophils are a "unique" population, it must be established to what extent they are unique. Immune cells cluster together on UMAPs, so what if these are a different cell type entirely, like another immature myeloid lineage, and this is an artifact of clustering? This could be clarified with a trajectory analysis and further subsetting of the immune population.

We thank the reviewers for this comment. We now realize that analyzing the cell cycle scores was not serving the intended purpose in this work. Moreover, due to the use of pooled samples for scRNA seq analyses, it may not be best to perform such downstream analyses in our datasets. We have thus removed these graphs from the revised version and have tried to simplify the conclusions of our study to the readers.

Our main take home from this study is the increase in number of mature (Ly6G+) and immature (Ly6G-) neutrophils in tobacco-flavored e-cig aerosol exposed mouse lungs as compared to air control. This result was validated using co-immunofluorescence in the revised manuscript (Figure 4).

In vivo validation of findings should be included, especially for the claimed changes. As of now, this paper serves more as a dataset that could be further explored by other groups, which in itself is valuable, but it is just one single cell sequencing experiment without validation.

We thank the reviewers for this comment. We have used multiple techniques (flow cytometry, multiplex protein assay, co-immunofluorescence) in the revised manuscript to validate the scRNA seq findings. However, this was a preliminary study which was designed to generate a small dataset for future experiments, and we do not have resources to add more validatory experiments for this study. We are currently designing chronic e-cig exposure studies to elaborate upon certain hypothesis generated through this study in future.

Minor Comments

There are several examples of typos or small errors in the text that would benefit from proofreading. Examples: line 51 "in the many countries including (the) United States (US), (the) United Kingdom..."; on line 54, the reference cited states that 9.4% of middle schoolers are daily users, not 9.2%; on line 55 the reference cited states that these are the most commonly used flavors, not the most preferred, which explains why the percentages do not add up to 100; line 120 "the lungs were in a collapsed state than the other groups"; line 127 "to confirm out speculations"; line 136 "PGVG" instead of the previously used "PG:VG"; line 140 "(single cell capture))"; line 999 "result in" rather than "results in" for Figure 4 title, etc.

We thank the reviewer for this comment. The manuscript has been thoroughly proofread and edited to avoid typos and grammatical errors.

If this is a "pilot study" (as it is stated in the introduction) it is meant to assess the validity of experimental design on a small scale to later test a hypothesis. The authors should change the phrasing.

We have now changed the phrasing as suggested.

The introduction lacked the necessary context and background. Some information described in the results section could be addressed in the intro. For example: What is the significance of neutrophils having a Ly6G deficiency? Why was the exposure duration of 1 hour a day for 5 days chosen? Why use nose-only exposure when many models use whole-body exposure? Why look at cell-type-specific changes?

We have made the necessary amendments in the introduction.

Some figure titles only address certain panels rather than summarizing the figure as a whole. For example, the title of Figure 1 only refers to panel D and is unrelated to serum cotinine levels, septa thickening, or mean linear intercept. The text discussed conclusions about septa thickening and Lm values for the fruit-flavored treatment group, so they are equally relevant to the figure compared to the metal levels.

We have now changed the Figures and Figure legends to summarize the figure.

significance level is not defined in Figure 1 legend although it is used in Figure 1C.

The Figure legend has now been updated.

Figure 1E does not include a scale bar.

We have now included the scale bar in updated figures.

The multiplex ELISA shown in the experimental design schematic is not further discussed in the paper. Flow cytometry plots should be displayed in addition to the data they generated.

The flow cytometry plots have now been included (Figures 3&5) and the results for Multiplex ELISA are shown as Figure S3D and lines 327-342 of the revised manuscript.

In Figure 1F, a multivariate ANOVA should be used so that multiple groups can be compared across sex, rather than plotting in a sex-specific manner and claiming there exists a sex bias. The small sample size also introduces an issue because a p-value cannot be generated with so few samples.

Per the suggestions made previously, figure 1F has now been removed from the revised manuscript.

The protocol for achieving a single-cell suspension should be detailed in the methods section. As is, it only describes the sample collection and preparation. This could help elucidate to the reader why the UMAP shows such a large abundance of immune cells.

We have now included the protocol in the revised manuscript.

Clarify whether PG:VG was used as a control in the scRNA sequencing in addition to air to generate the UMAP in Figure 2A.

Yes, PG:VG was used as one of the controls which has now been illustrated as groupwise comparison in Figure 2D. We have also included the comparisons to identify DEGs in myeloid and lymphoid clusters upon comparison of various treatment groups versus PGVG (Supplementary Files S13-S18)

A UMAP should be shown for each treatment group/flavor. The overall UMAP in Figure 1A is good, but there could be another panel with separate projections for each condition.

A groupwise UMAP has now been included in Figure 2D.

In Figure 2C, relative cell percentage is not a reliable method to quantify cell type and the histogram is not a great way to visualize the data or its statistical significance. These claims should also be validated in tissue.

We thank the reviewers for this comment and have tried to validate the findings using Flow cytometry. However, we may want to add that the changes observed in single cell technologies cannot be validated using simple molecular biology techniques as the markers used to specify cell clusters in scRNA seq is too specific which was not the case for the design of flow panel in this work. Our major purpose of using cell percentages was to show the flavor-specific changes in generalized cell populations in mouse lungs. So, we have still included these graphs in the revised manuscript.

Figure 2D could be better illustrated with a volcano plot to show which genes are being dysregulated rather than just how many. Knowing which genes are affected is more valuable than knowing just the number of genes.

Figure 2D is no longer a part of the revised manuscript. For the other comparisons we have still used heatmaps as they also depict sex-specific changes in gene expressions, which would have been difficult to elucidate using volcano plots.

Assuming Figure 3C is representative of all conditions, then Figures 3C and D demonstrate that Ly6G- neutrophils are present in all conditions including controls. To see whether they are truly present in different abundances between treatment and control groups, separate UMAPs of the neutrophil subsets should be made per condition or use a dot plot for Figure 3A. This also applies to Figure 3B.

We thank the reviewers for pointing this out. We have now revamped the whole manuscript and used additional validation experiments to show the presence of Ly6G- and Ly6G+ neutrophil population upon exposure to tobacco-flavored e-cig aerosols.

Figure 3E shows that there is no statistically significant change in % of Ly6G+ neutrophils across treatment groups, but the text claims that there is "an increase in the levels of Ly6G+ neutrophils in lung digests from mouse lungs exposed to tobacco-flavored e-cig aerosols" (lines 207-209). The text also claims that "The observed increase was more pronounced in males as compared to females" (lines 209-210), but there was no statistical analysis across sexes to support this statement. It is clear that the change in % of Ly6G+ neutrophils is more pronounced in males than females, but it is still not statistically significant. This figure should also be repeated for analysis of Ly6G- neutrophils. Lines 272-274 mention that the % increase is higher for Ly6G- neutrophils than for Ly6G+ neutrophils, but there is not an analogous histogram to demonstrate this. The claims made in lines 275-280 are not clearly shown in any figure.

We thank the reviewer for this query. This was an error on our part. We have now added sex-specific changes using scRNA seq, flow cytometry and co-immunofluorescence-based experiments to prove that more pronounces changes in the Ly6G+ and Ly6G- neutrophil population occurs in female mice and not males.

Figures 4 and 6 have an overwhelming amount of heatmaps. Volcano plots with downstream analyses could be used to make some of this data more legible. The main findings should be validated in vivo/in tissue.

We have now revamped the figures and data distribution to make the data legible and remove overwhelming amount of data from the slides.

For Figure 5, show cell type by condition and do differential gene expression analysis displayed in a volcano plot. Then, stain tissue to validate the findings. Compare across sex during statistical analysis.

The necessary changes have been made.

Figure 6 error: panels E and F should be labeled as "tobacco" rather than "fruit".

Error has now been fixed.

Figure 7C can be placed in the supplemental materials.

It has now been included in supplemental materials.

The Figure 6E title should have been tobacco instead of fruit.

This error has now been fixed.

Line 381 mentioned the wrong subfigure. (Figure 7B instead of 7E).

We have now made the necessary edits.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: E-cigarettes (e-cigs) are a public health concern for young adults due to their popularity and evidence for increased oxidative stress and immunotoxicity. Yet an extensive study defining the cell-specific immune changes upon exposure to flavored e-cigs remains elusive. To understand the immunological lung landscape upon acute nose-only exposure of C57BL/6J to flavored e-cig aerosols we performed single-cell RNA sequencing (scRNA seq). scRNA profiles of 71,725 cells were generated from control and treatment groups (n=2/sex/group). A distinct phenotype of Ly6G-neutrophils was identified in lungs exposed to tobacco flavored e-cig aerosol which demonstrated dampened IL-1 mediated and pattern recognition signaling as compared to air controls. Differential gene expression analyses identified dysregulation of T-cell mediated pro-inflammation (Cct7, Cct8) and stress-response signals (Neurl3, Stap1, Cirbp and Htr2c) in the lungs of mice exposed to e-cig aerosols, with pronounced effects for tobacco flavor. Flow cytometry analyses and cytokine/chemokine assessments within the lungs corroborated the scRNA seq data, demonstrating a significant increase in T-cell percentages and levels of T-cell associated cytokine/chemokines in the lungs of tobacco-flavored aerosol exposed mice. Increased levels of Klra4 and Klra8 expression also suggest an enhanced natural killer (NK) cell activity in this mouse group. Overall, this is a pilot study identifying increase in the percentages of Ly6G-neutrophils that may be responsible for dampened innate immune responses and heightened T-cell cytotoxicity in lungs of tobacco-flavored e-cig aerosol exposed mice. In addition, we provide preliminary evidence for sex-specific changes in the transcriptional landscape of mouse lungs upon exposure to e-cig aerosol, an area that warrants further study.

Revised version: E-cigarettes (e-cigs) are a public health concern for young adults due to their rising popularity despite evidence of harmful effects, i.e. increased oxidative stress and immunotoxicity. Yet an extensive study defining the cell-specific immune changes upon exposure to flavored e-cigs remains elusive. To determine the immunological lung landscape upon acute nose-only exposure of C57BL/6J to flavored e-cig aerosols, we performed single-cell RNA sequencing (scRNA seq). scRNA profiles of 71,725 cells were generated from control and treatment groups (n=2/sex/group). An increase in the Ly6G- neutrophil population was identified in lungs exposed to tobacco flavored e-cig aerosol. Differential gene expression analyses identified dysregulation of T-cell mediated pro-inflammation (Cct7, Cct8) and stress-response signals (Neurl3, Stap1, Cirbp and Htr2c) in the lungs of mice exposed to e-cig aerosols, with pronounced effects for tobacco flavor. Flow cytometry analyses and cytokine/chemokine assessments within the lungs corroborate the scRNA seq data, demonstrating an increase in the T-cell percentages and levels of T-cell associated cytokine/chemokines in the lungs of tobacco-flavored aerosol exposed mice. Increased levels of Klra4 and Klra8 expression also suggest an enhanced natural killer (NK) cell activity in this mouse group. Interestingly, the increase in the level of Ly6G – (immature) neutrophils upon exposure to tobacco-flavored e-cig aerosols was noticed for female C57BlL6J mice and not their male counterparts which was validated through immunofluorescence staining. Overall, this study identifies sex-specific increase in the percentages of Ly6G- neutrophils that may be responsible for dampened innate immune responses in females and heightened T-cell cytotoxicity in male mouse lungs of tobacco-flavored e-cig aerosol exposed mice.

Introduction

Earlier version: Electronic cigarettes (e-cigs) or electronic nicotine delivery systems are a relatively novel set of tobacco and flavored products that have gained immense popularity amongst adolescents and young adults in the many countries including United States (US), United Kingdom (UK) and China. Flavors are one of the key features that make these products alluring to the younger diaspora (1). Reports indicate that in 2020 about 22.5% of high school students and 9.2% of middle school students in the US were daily vapers or e-cig users with fruit (66%), mint (57.5%), and menthol (44.5%) being the most preferred flavors (2). However, not much is known about the flavor-specific effects of e-cig vaping on the health and immunity of an individual.

Revised version: Electronic cigarettes (e-cigs) or electronic nicotine delivery systems are a relatively novel set of tobacco/nicotine and flavored products that have gained immense popularity amongst adolescents and young adults in many countries including the United States (US), the United Kingdom (UK) and China. Flavors are one of the key features that make these products alluring to the younger diaspora (1). Reports indicate that in 2020 about 22.5% of high school students and 9.4% of middle school students in the US were daily vapers or e-cig users with fruit (66%), mint (57.5%), and menthol (44.5%) being the most commonly used flavors (2). However, not much is known about the flavor-specific effects of e-cig vaping on the health and immunity of an individual specially focusing on cell-types and gene transcripts.

Earlier version: E-cig products and aerosols are known to contain harmful constituents including formaldehyde, benzaldehyde, acrolein, n-nitrosamines, volatile organic compounds (VOCs), ketenes, and metal ions (3–5). Studies have indicated that exposure to e-cigs may enhance inflammatory responses, oxidative stress, and genomic instability in exposed cells or animal systems (6–8). Risk assessment of the non-carcinogenic (systemic) risks of inhaled diacetyl, a potential component of e-liquids, has estimated the hazard quotient to be greater than 1 amongst teens (9). Furthermore, clinical and in vivo studies have suggested that exposure to e-cig aerosols could impair innate immune responses in the host thus making them more susceptible to bacterial/viral infections. The bacterial clearance, mucous production, and phagocytic responses in these individuals are shown to be affected upon use of e-cigs (10–14).

Revised version: E-cig products and aerosols are known to contain harmful constituents including formaldehyde, benzaldehyde, acrolein, n-nitrosamines, volatile organic compounds (VOCs), ketenes, and metal ions (3–5). Studies have indicated that exposure to e-cigs may enhance inflammatory responses, oxidative stress, and genomic instability in exposed cells or animal systems (6–8). Risk assessment (systemic) of inhaled diacetyl, a potential component of e-liquids, has estimated the non-carcinogenic hazard quotient to be greater than 1 amongst teens (9). Furthermore, clinical and in vivo studies have suggested that exposure to e-cig aerosols could impair innate immune responses in the host thus making them more susceptible to bacterial/viral infections. The bacterial clearance, mucous production, and phagocytic responses in these individuals are shown to be affected upon use of e-cigs (10–14).

Earlier version: However, cell-specific changes within the lung upon vaping is not fully understood making it hard to determine the long-term health impacts of the use of these novel products. In this pilot study, we aim to determine effects of e-cig exposure on mouse lungs at single cell level. To do so, we exposed C57BL/6J mice to 5-day nose-only exposure to air, propylene glycol: vegetable glycerin (PG:VG), fruit-, menthol- and tobacco-flavored e-cig aerosols. We performed single-cell RNA sequencing (scRNA seq) on the lung digests from exposed and control animals and identified neutrophils and macrophages, among others, as the major cell populations in the lung that were affected upon acute exposure. We were able to identify 29 gene targets that were commonly dysregulated amongst all our treatment groups upon aggregating results from the major lung cell types identified in this study. These gene targets are the markers of early immune dysfunction upon e-cig aerosol exposure in vivo and could be studied in detail to understand temporal changes in their expression that may govern allergic responses and adverse pulmonary health outcomes upon acute and sub-acute exposures to e-cigarette aerosols.

Revised version: However, cell-specific changes within the lung upon vaping is not fully understood, making it hard to determine the long-term health impacts of the use of these novel products. In this respect, single cell technology is a powerful tool to analyze gene expression changes within cell populations to study cell heterogeneity and function (15–17). Such an investigation is important to deduce the health effects of acute and chronic use if e-cigs in young adults. In this study, we aim to determine effects of e-cig exposure on mouse lungs at single cell level. To do so, we exposed C57BL/6J mice to 5-day nose- only exposure to air, propylene glycol: vegetable glycerin (PG:VG), fruit-, menthol- and tobacco-flavored e-cig aerosols. The nose-only exposure has more translational relevance over whole-body exposure (18), owing to which, we chose nose-only exposure profile for this work. To limit the stress to the animals a 1-hour exposure was chosen per day. We performed single-cell RNA sequencing (scRNA seq) on the lung digests from exposed and control animals and identified neutrophils and macrophages, among others, as the major cell populations in the lung that were affected upon acute exposure. We were able to identify 29 gene targets that were commonly dysregulated amongst all our treatment groups upon aggregating results from the major lung cell types identified in this study. These gene targets are the markers of early immune dysfunction upon e-cig aerosol exposure in vivo and could be studied in detail to understand temporal changes in their expression that may govern allergic responses and adverse pulmonary health outcomes upon acute and sub-acute exposures to e-cigarette aerosols.

Results

Added in the revised version: Exposure to e-cig aerosol results in flavor-dependent exposure to different metals and mild histological changes in vivo

Removed in the revised version: Fruit-flavored e-cig aerosol results in lung injurious responses upon acute exposure in vivo

Earlier version: This study was designed to characterize the effects of exposure to flavored e-cig aerosol at single-cell level to understand the immunological changes in the lung microenvironment upon e-cig use. To do so, we generated a first ever single-cell profile of e-cig aerosol exposed mouse lungs (n= 2/sex/group). The thus obtained results were then validated with the help of our validation cohort of n= 3/sex/group as shown in Fig 1A. Since all the commercially available e-liquids used in this study contained tobacco derived nicotine, we first determined the level of serum cotinine (a metabolite of nicotine) to prove successful exposure to the mice in each treatment group. As expected, we did not see any traces of cotinine in the serum of air and PG:VG exposed mice. Significant levels of cotinine were detected in the serum of mice exposed to fruit-, menthol-, and tobacco-flavored e-cig aerosols (Fig 1C).

Revised version: This study was designed to characterize the effects of exposure to flavored e-cig aerosols at single-cell level to understand the immunological changes in the lung microenvironment upon e-cig use. To do so, we generated a first ever single-cell profile of e-cig aerosol exposed mouse lungs (n= 2/sex/group). The thus obtained results were then validated with the help of our validation cohort of n= 3/sex/group as shown in Fig 1A. Since all the commercially available e-liquids used in this study contained tobacco derived nicotine, we first determined the level of serum cotinine (a metabolite of nicotine) to prove successful exposure to the mice in each treatment group. As expected, we did not see any traces of cotinine in the serum of air and PG:VG exposed mice. Significant levels of cotinine were detected in the serum of mice exposed to fruit-, menthol-, and tobacco-flavored e-cig aerosols (Figure S1B).

Earlier version: E-cig aerosols show flavor-dependent changes in the levels of quantified metals

Revised version: Flavor-dependent changes in the levels of quantified metals, but no major histological damage on acute exposure to flavored e-cig aerosol in mice.

Earlier version: Schematics showing the exposure profile and experimental design to understand the effect of exposure to differently flavored e-cig aerosols in the lungs from C57BL/6J mice using scRNA seq (A) and the nose only exposure system used for performing the mouse experiment (B). The exposure characteristics were assessed by measuring the serum cotinine levels in the blood of exposed and control mice. Data are shown as mean ± SEM (n = 4/group); ns: not significant. **p<0.01, and ***p<0.001 vs air, per one-way ANOVA for multiple comparison (C) and analyzing the levels of metals in the aerosols captured from each day of exposure using ICP-MS (D). Lung morphometric changes as observed using H&E staining of lung slices from air, PGVG and differently flavored e-cig exposed mice. Representative images of n = 2/sex/group at 10X magnification is provided (E) and quantified values for mean linear intercept (Lm) plotted in a sex-specific manner (F)

Revised version: Schematics showing the exposure profile and experimental design to understand the effect of exposure to differently flavored e-cig aerosols in the lungs from C57BL/6J mice using scRNA seq (A)). Bar Graph showing the levels of metals in the aerosols captured each day of exposure using ICP-MS (B). Lung morphometric changes as observed using H&E staining of lung slices from air, PG:VG and differently flavored e-cig exposed mice. Representative images of n = 2/sex/group at 10X magnification is provided (C).

Earlier version: Since metals released upon heating of the coils of e-cig devices are a source of toxicity upon vaping (15, 16), we further monitored the levels of metals in the e-cig aerosols generated during each day of mouse exposures. This acted as an indirect measure of characterizing the chemical properties of the aerosols used for exposure in this study. To monitor the release of metals into the lungs of the animals, the condensate from each day of exposure was collected and the levels of select elements were detected using ICP-MS. A detailed account of the concentrations of identified elements/metals is provided in Table 1. Interestingly, we identified flavor-dependent changes in the levels of metals like Ni, Zn, Na, K and Cu. Of note, despite the use of same wattage and temperature (max of 230⁰C) for generation of e-cig aerosols, the leaching of each metal varied per day of exposure (Fig 1D). This is a crucial result as it highlights the importance of studying the impacts of coil composition and device design on the chemical composition of generated e-cig aerosols. These variations might affect the risk and toxicity associated with each of these products, an area that warrants further study. Next, we performed H&E staining of the lung tissue sections to study the morphometric changes in the mouse lungs upon exposure to differently flavored e-cig aerosols. We did not find much evidence of tissue damage or airspace enlargement upon acute exposures in our model as expected. However, we found evidence of increased alveolar septa thickening in the lungs of both male and female mice exposed to fruit-flavored e-cig aerosols (Figure 1E). However, this could be a result of lower agarose inflation observed for the mouse lungs in this group of mice. We had challenges/difficulties with inflating the mouse lungs for all the mice in this group, however male mice showed more restriction than the female. Owing to the lack of proper inflation, the lungs were in a collapsed state than the other groups where the lungs were fully inflated. We have also performed mean linear intercept (Lm) measurements for all the lung tissue sections to show lowered Lm values for mouse lungs exposed to fruit-flavored e-cig aerosols as compared to air controls (Figure 1F). Airway obstruction could be indicative of “atelectasis”, a condition that may result from lack of surfactant or alveolar damage thus affecting the lung expansion in these mice. Since this was not the prime focus of this study, we did not conduct further experiments to confirm out speculations, but nevertheless, we report these changes to inform and encourage further research in this area.

Revised version: Since metals released upon heating of the coils of e-cig devices are a source of toxicity upon vaping (19, 20), we further monitored the levels of metals in the e-cig aerosols generated during each day of mouse exposures. This acted as an indirect measure of characterizing the chemical properties of the aerosols used for exposure in this study. To monitor the release of metals into the lungs of the animals, the aerosol condensate from each day of exposure was collected and the levels of select elements were detected using ICP-MS. A detailed account of the concentrations of identified elements/metals is provided in Table 1. Interestingly, we identified flavor-dependent changes in the levels of metals like Ni, Zn, Na, K and Cu. Note, despite the use of same wattage and temperature (max of 230⁰C) for generation of e-cig aerosols, the leaching of each metal varied per day of exposure (Figure 1B). This is a crucial result as it highlights the importance of studying the impacts of coil composition and device design on the chemical composition of generated e-cig aerosols. These variations might affect the risk and toxicity associated with each of these products, an area that warrants further study.

Added in the revised version: Next, we performed H&E staining of the lung tissue sections to study the morphometric changes in the mouse lungs upon exposure to differently flavored e-cig aerosols. We did not find much evidence of tissue damage or airspace enlargement upon acute exposures in our model, as expected. However, we found evidence of increased alveolar septa thickening in the lungs of both male and female mice exposed to fruit-flavored e-cig aerosols (Figure 1C). However, this could be a result of lower agarose inflation observed for the mouse lungs in this group. We had challenges/difficulties with inflating the mouse lungs for all mice in this group, however male mice showed more restriction than the female. Owing to the lack of proper inflation, the lungs were in a collapsed state than the other groups which could be a probable cause of the histological observations. Airway obstruction could be indicative of “atelectasis”, a condition that may result from lack of surfactant or alveolar damage thus affecting the lung expansion in these mice. Since this was not the prime focus of this study, we did not conduct further experiments to confirm our speculations, but nevertheless, we report these changes to inform and encourage further research in this area.

Detailed map of cellular composition during acute exposure to e-cig aerosols reveals distinct changes in immune cell phenotypes

Earlier version: As the principle focus of this study was to identify the flavor-dependent and independent toxicities upon exposure to commercially available e-cig aerosols at the single-cell level, we performed the scRNA seq on the mouse lungs from exposed and control mice. After quality control filtering (Figure S1), normalization and scaling, we generated scRNA seq profiles of 71,725 cells in total. Except for the PGVG group, all the rest of the treatments had approximately similar cell viabilities, cell capture, and other quality assessments in this study. However, for uniform analyses, equal features/ genes were used across all the groups for subsequent analyses. A detailed account of the cell viability, number of samples (single-cell capture)), and gene features identified before and after filtering upon QC check of sequenced data is provided in Table S1.

Revised version: As the principle focus of this study was to identify the flavor-dependent and independent toxicities upon exposure to commercially available e-cig aerosols at the single-cell level, we performed the scRNA seq on the mouse lungs from exposed and control mice. After quality control filtering (Figure S2), normalization and scaling, we generated scRNA seq profiles of 71,725 cells in total. Except for the PG:VG group, all the rest of the treatments had approximately similar cell viabilities, cell capture, and other quality assessments in this study. However, for uniform analyses, equal features/ genes were used across all the groups for subsequent analyses. A detailed account of the cell viability, number of samples (single-cell capture), and gene features identified before and after filtering upon QC check of sequenced data is provided in Supplementary File S1.

Earlier version: Uniform Manifold Approximation and Projection (UMAP) was used for dimensionality reduction and visualization of cell clusters. Cell annotations were performed based on the established cell markers in Tabula Muris database and available published literature. The general clustering of individual cell types based upon the commonly known cell markers was used to identify -Endothelial (identified by expression of Cldn5), Epithelial (identified by expression of Sfgtpa1), Stromal (identified by expression of Col3a1), and Immune (identified by expression of Ptprc) cells. The Immune cell population was further subclustered into lymphoid and myeloid populations in the UMAP analyses (Figure 2A&B). The relative cell percentages were plotted in a sex-dependent manner. Of note, upon assessing the sex- and treatment-specific changes in observed cell frequencies for individual cell types using two-way ANOVA, we did not find any significant variation in the relative frequencies of the general cell types in our treatment groups across both sexes (Figure 2C). A detailed account of the two-way ANOVA statistics for individual cell types is provided in Table S2.

Revised version: Uniform Manifold Approximation and Projection (UMAP) was used for dimensionality reduction and visualization of cell clusters. Cell annotations were performed based on the established cell markers in Tabula Muris database and available published literature and we identified 24 distinct cell clusters as shown in Figure 2A. The general clustering of individual cell types based upon the commonly known cell markers was used to identify -Endothelial (identified by expression of Cldn5), Epithelial (identified by expression of Sftpa1), Stromal (identified by expression of Col3a1), and Immune (identified by expression of Ptprc) cells (Figure 2B). The ‘FindVariableFeatures’ from Seurat was used to identify cell-to-cell variation between the identified clusters and the top 2000 variable genes identified in each cluster have been elaborated in Supplementary File S2. Upon plotting the number of counts/cells identified in each cluster in a sex- and treatment-specific manner we observed minor changes, though not significant, in the count of myeloid and lymphoid cell clusters in e-cig aerosol exposed mouse lungs as compared to air controls (Figure 2C). A detailed account of the two-way ANOVA statistics for individual cell types is provided in Supplementary File S3A.

Earlier version: scRNA seq analyses reveals maximum changes in the cell states of immune cell population upon exposure to differently flavored e-cig aerosols.

Revised version: scRNA seq analyses reveal maximum changes in the cell states of immune cell population upon exposure to differently flavored e-cig aerosols.

Earlier version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to differently flavored e-cig aerosols. The mice were sacrificed after the final exposure and lungs from air (control) and differently flavored e-cig aerosol (fruit, menthol, and tobacco)-exposed mice were used to perform scRNA seq. UMAP plot of 71,725 cells captured during scRNA seq showing the four major cell populations identified from control and experimental mouse lungs (A) and the expression of canonical markers used for identifying stromal (Col3a1), epithelial (Sftpa1), endothelial (Cldn5) and immune (Ptprc) cell populations. The intensity of expression is indicated by the red-yellow coloring (B). The cell frequencies (plotted as percentage) of different cell clusters in each sample type showing the sex-dependent variations in the cell composition on exposure to differently flavored e-cig aerosols (C). Bar graphs showing the number of significant DEGs (p <0.01) found to be up (blue) and down (orange)-regulated in each cell cluster in differently flavored e-cig aerosol exposed mouse lungs as compared to air controls as determined by aggregate analyses for both sexes in each group using DESeq2 (D).

Revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to differently flavored e-cig aerosols. The mice were sacrificed after the final exposure and lungs from air (control) and differently flavored e-cig aerosol (fruit, menthol, and tobacco)-exposed mice were used to perform scRNA seq. UMAP plot of 71,725 cells captured during scRNA seq showing the four major cell populations identified from control and experimental mouse lungs (A) and the expression of canonical markers used for identifying stromal (Col3a1), epithelial (Sftpa1), endothelial (Cldn5) and immune (Ptprc) cell populations. The intensity of expression is indicated by the red-yellow coloring (B). The cell frequencies (plotted as counts) of different cell clusters in each sample type showing the sex-dependent variations in the cell composition on exposure to differently flavored e-cig aerosols (C). Group-wise comparison of the UMAPs upon comparing PG:VG (blue), fruit (yellow), menthol (green) and tobacco (red) versus air (air) groups following dimensionality reduction and clustering of scRNA seq data (D). Here, AT1: alveolar type I, AT2: alveolar type II, SMC: smooth muscle cell, gCap: general capillary, aCap: alveolar capillary, and NK: natural killer.

Earlier version: DESeq2 was then used to identify the differentially regulated genes in each broad cell category. A list of significant (padj<0.05) differentially expressed genes (DEGs) was generated using aggregate analysis for both sexes for each treatment group and number of up- and down-regulated genes were plotted in a flavor-dependent manner to get a snapshot of the cell-types and flavor-categories that show most dysregulation upon exposure (Figure 2D). Our data showed dysregulation of genes in endothelial, epithelial, and stromal cell populations, but maximum effect was observed on the cells from immune cell clusters. This is not surprising, as the immune system especially the myeloid cells form the frontline of the host’s defenses against external stressors (17, 18). Compared to air, we observed a dysregulation of 553 genes (338 upregulated; 215 downregulated) in the myeloid cell cluster of mouse lungs exposed to tobacco-flavored e-cig aerosol. We identified 324 and 24 DEGs in the myeloid lung cluster from mouse lungs exposed to menthol- and fruit-flavored e-cig aerosols respectively as compared to air control. For the lymphoid cluster, we observed maximum dysregulation in the lungs exposed to fruit-flavored e-cig aerosols with a total of 112 DEGs. In contrast, 41 and 9 significant DEGs were identified for lymphoid cluster from lungs exposed to tobacco and menthol-flavored aerosols respectively.

Revised version: Groupwise comparisons of each treatment group versus Air control, did not show many changes in the cell clusters thus showcasing little to no effect on the overall lung cellular compositions in treated and control groups. However, the number of cells identified in each cluster for PG:VG treatment group was very low. We would like to report that this is an outcome of the low viability of this sample prior to scRNA seq, and not necessarily the effect of the treatment (Figure 2D). Differential gene expression analyses showed dysregulation of genes in all cell populations, but maximum effect was observed on the cells from immune cell clusters. This is not surprising, as the immune system, especially the myeloid cells form the frontline of the host’s defenses against external stressors (21, 22). Compared to air, we observed a dysregulation of 553 genes (338 upregulated; 215 downregulated) in the myeloid cell cluster of mouse lungs exposed to tobacco-flavored e-cig aerosol. We identified 324 and 24 DEGs in the myeloid lung cluster from mouse lungs exposed to menthol- and fruit-flavored e-cig aerosols respectively as compared to air control (Supplementary Files S5-S7). For the lymphoid cluster, we observed maximum dysregulation in the lungs exposed to fruit-flavored e-cig aerosols with a total of 112 DEGs. In contrast, 41 and 9 significant DEGs were identified for lymphoid cluster from lungs exposed to tobacco and menthol-flavored aerosols respectively (Supplementary Files S9-S11).

Earlier version: It is important to mention here that most of the commercially available e-liquids /e-cig products use PG:VG as the base liquid to generate the aerosol and act as a carrier for flavoring chemicals. Thus, we compared the effect of PG:VG alone in our study to make further comparisons between individual flavoring products with that of PG:VG only. DESeq2 analyses showed dysregulation of 276 genes in mouse lungs exposed to PG:VG alone in the myeloid cell cluster as compared to air controls. Contrary to this, exposure to PG:VG aerosols affected 24 genes in the lymphoid cluster as compared to air control. Furthermore, when compared to PG:VG, exposure to fruit-flavored e-cig aerosol dysregulated 262 genes in the myeloid cluster and 37 genes in the lymphoid cluster (Table 2).

Revised version: It is important to mention here that most of the commercially available e-liquids /e-cig products use PG:VG as the base liquid to generate the aerosol and act as a carrier for flavoring chemicals. Thus, we compared the effect of PG:VG alone in our study to make further comparisons between individual flavoring products with that of PG:VG only. DESeq2 analyses showed dysregulation of 276 genes in mouse lungs exposed to PG:VG alone in the myeloid cell cluster as compared to air controls (Supplementary Files S4). Contrary to this, exposure to PG:VG aerosols affected 24 genes in the lymphoid cluster as compared to air control (Supplementary Files S8). Furthermore, when compared to PG:VG, exposure to fruit, menthol and tobacco-flavored e-cig aerosol dysregulated 262, 873 and 960 genes in the myeloid cluster and 37, 64 and 112 genes in the lymphoid cluster respectively. A detailed account of the DEGs identified upon comparing fruit, menthol and tobacco-flavored e-cig aerosol exposed mouse lungs to PG:VG in myeloid and lymphoid clusters have been provided in Supplementary Files S13-S18.

Added in the revised version: Exposure to differently flavored e-cig aerosols results in sex-specific variations in neutrophilic and eosinophilic response in the mouse lungs

Added in the revised version: We found dysregulation in immune cell population, more specifically the myeloid cells, upon exposure to e-cig aerosols through our previous analyses. We thus plotted the changes in cell counts across the various treatment groups and control in a sex-dependent manner. scRNA seq results did not show any change in the cell frequencies of alveolar macrophages in flavored e-cig aerosol exposed mouse lungs as compared to controls (Figure 3A; Supplementary File S3B). However, the cell count of neutrophils in tobacco-flavored e-cig aerosol exposed mice lungs were found to be increased in females, but not in males (Figure 3B, Supplementary File S3B). Taken together, we show a significant shift towards the neutrophil-mediated immune response in the lungs of female mice exposed to tobacco-flavored e-cig aerosols. We used flow cytometry to validate the scRNA seq changes. While the changes observed in the levels of alveolar macrophages in the mouse lungs corroborated with our scRNA seq findings, the results for the neutrophil cell cluster did not show matching trend. This could be due to the increase in the levels of immature neutrophils in the lungs of female tobacco-e-cig exposed mouse lungs characterized by the absence of Ly6G cell surface marker, which was used to stain for neutrophils via flow cytometry (Figure 3C-D, Supplementary File S3D).

Removed in the revised version: Exposure to tobacco-flavored e-cig aerosols dampens IL-1 mediated signaling in Ly6G-neutrophil cluster

Removed in the revised version: We found maximum dysregulation in immune cell population, more specifically the myeloid cells, upon exposure to e-cig aerosols through our previous analyses. We thus subclustered the myeloid and lymphoid cell populations to further sub-divide their cell types and identify the flavor-dependent changes in immune cell responses. Upon sub clustering the myeloid cell population, we identified 14 unique clusters comprising all the major cell phenotypes including neutrophils, alveolar macrophages (AM), interstitial macrophages (IM), monocytes, dendritic cells (DC), and mast cells (Figure 3A). A detailed account of all the cell types identified with their respective marker genes is provided in Table 3. Cell cycle scoring revealed the neutrophil cluster to be in G2M phase (Figure 3B), indicating that these cell types are either at a checkpoint for uncontrolled proliferation or advancing toward a state of exhaustion (19, 20). In either case, we believed that studying these cells might be of interest to understand the mechanisms governing the immune dysregulation upon exposure of C57BL/6J mice to differently flavored e-cig aerosols. On deeper evaluation, we identified two unique phenotypes of neutrophils (identified by S100a8, Cxcl2, Sell and Lcn2) in the mouse lungs. These clusters were named as Ly6G+ Neutrophils and Ly6G− Neutrophils based on the presence or absence of Ly6G marker, respectively (Figure 3C). Ly6G is important for neutrophil migration and is crucial to regulate neutrophil function within the lung (21). Interestingly we found a consistent flavor-dependent increase (though not significant) in the Ly6G− neutrophil percentage in the lungs of exposed mice as compared to air controls. To validate the changes in cell percentages upon flavor exposure from scRNA seq, we performed flow cytometry. Our data from flow cytometry corroborated with the scRNA analyses results showing an increase in the levels of Ly6G+ neutrophils in lung digests from mouse lungs exposed to tobacco-flavored e-cig aerosols. The observed increase was more pronounced in males as compared to females thus highlighting the sex-dependent nature of changes in immune cell frequencies within the exposed lungs (Figure 3E).

Earlier version: Cellular composition of myeloid cells in air and e-cig aerosol exposed mouse lung reveal a unique population of neutrophils.

Revised version: Cellular composition of myeloid cells in air and e-cig aerosol exposed mouse lung reveal sex-specific increase in neutrophil count through scRNA seq but not flow cytometry.

Added in the revised version: Cell frequencies (represented as counts) of alveolar macrophages (A) and neutrophils (B) across controls and flavored e-cig aerosol exposed mouse lungs were plotted in a sex-specific manner. Representative flow plots (C) and bar graphs (D) showing sex-dependent changes in the percentages of neutrophils (CD45+ CD11b+ Ly6G+) and alveolar macrophages (CD45+ CD11b-SiglecF+) populations in lung digests from mice exposed to differently flavored e-cig aerosols. Data are shown as mean ± SEM (n = 3/sex/group).

Added in the revised version: It is important to mention that due to the presence of nicotine in all the e-liquids used as treatment groups for this study, we added an extra control group- PG:VG+Nic- for select experiments. Flow cytometric analyses revealed slight variations in the lung neutrophils and macrophage percentages observed in PG:VG+Nic group as compared to PG:VG only. However, none of these changes were significant. Furthermore, the patterns of change observed for the lungs exposed to aerosols from flavored e-liquids were quite distinct from those observed for PG:VG+Nic thus proving that the observed changes are not solely due to the presence of nicotine in these groups (Figures S3A-B). To probe further, we subclustered the myeloid cell populations. Upon sub clustering, we identified 14 unique clusters comprising all the major cell phenotypes including neutrophils, alveolar macrophages (AM), interstitial macrophages (IM), monocytes, dendritic cells (DC), and mast cells (Figure S3C). A detailed account of all the cell types identified with their respective marker genes is provided in Supplementary File 12. On deeper evaluation, we identified two unique phenotypes of neutrophils (identified by S100a8, Cxcl2, Sell and Lcn2) in the mouse lungs. These clusters were named as Ly6G+ Neutrophils and Ly6G- Neutrophils based on the presence or absence of Ly6G marker, respectively (Figure S3C). Ly6G is important for neutrophil migration, maturation and function within the lung (23). Interestingly we found an increase in the levels of both Ly6G+ and Ly6G- neutrophil in the tobacco-flavored e-cig aerosol exposed female mouse lungs as compared to air control. Though we did not observe similar changes in male mice exposed to tobacco-flavored e-cig aerosol (Figure 4A-B, Supplementary File S3C). The scRNA seq findings were further validated using co-immunofluorescence using S100A8 (red, pan-neutrophil marker) and Ly6G (green, marker of mature neutrophil). Co-immunofluorescence results show increase in levels of both Ly6G + (yellow stained) and Ly6G – (red stained) neutrophil in female mice exposed to tobacco flavored e-cig aerosol, but not males further corroborating our previous findings and suggesting that sex-specific changes in innate immune response in vivo (Figures 4C-D).

Added in the revised version: Co-immunofluorescence validates the increase of Ly6G– and Ly6G+ neutrophil population in the tobacco flavored e-cig exposed female C57BL/6J mice.

Added in the revised version: The myeloid cell clusters were subsetted to identify two populations of neutrophils with and without the presence of Ly6G cell marker representing mature and immature neutrophils respectively (Figure S3C). Bar graph showing the sex-specific changes in the cell frequencies (denoted as counts) of Ly6G+ (A) and Ly6G- (B) neutrophils in the lung of tobacco-flavored e-cig aerosol exposed mouse lung as compared to air control. scRNA seq findings were validated by staining the tissue sections from tobacco-flavored e-cig aerosols and control (air) with Ly6G (green) and S100A8 (red, pan-neutrophil marker). Representative images showing the co-immunostaining of Ly6G and S100A8 (shown as yellow puncta) in female (C) and male (D) control and tobacco-flavored e-cig aerosol exposed mice. Here Ly6G+ neutrophils are represented by double staining using yellow arrows, while Ly6G- neutrophils are denoted by red arrows.

Added in the revised version: Though we did not find a distinct eosinophil cluster through our scRNA seq analyses, a significant decline in the level of eosinophils in the lungs of female mice exposed to menthol-flavored e-cig aerosol was observed through flow cytometric analyses. Though not significant, marked decline in the eosinophil count was observed upon acute exposure of male and female C57BL/6J mice to tobacco-flavored e-cig aerosols. Contrarily, the eosinophil count in the lung of male mice exposed to fruit flavored e-cig aerosols was observed to be significantly higher as compared to PG:VG exposed lungs (Figures S4).

Removed in the revised version: A total of 14 unique cell clusters were identified from the myeloid cell subset of air and e-cig aerosol exposed mouse lung (A). UMAP plot showing the cell cycle state of different cell types in the myeloid cluster as determined using the Cell cycle scoring in Seurat (B). The heatmap showing the most DEGs in each cluster across the myeloid subset. These differential expressions were used to determine and name the cell types in this cluster. The intensity of expression is indicated by the purple coloring (C). The cell percentages of different cell types in the myeloid cluster showing the sex-dependent variations in the cell composition on exposure to differently flavored e-cig aerosols (D). Flow cytometry results showing sex-dependent changes in the percentages of lung neutrophils (CD45+ CD11b+ Ly6G+) in lung tissue digest (E) and alveolar macrophages (CD45+ CD11b-SiglecF+) in BALF (F) from mice exposed to differently flavored e-cig aerosols. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, and ****p<0.0001, per Tukey post hoc two-way ANOVA comparison.

Removed in the revised version: AM form the bulk of the myeloid cell population in the lung and are characterized by top markers like Chil3, Plet1, Lpl and Marco (Figure 3C) (22). A flavor-dependent decrease in the resident AM (RAM) population was observed for both menthol and tobacco flavored e-cig aerosol exposed lungs (Figure 3D). These changes were again corroborated through the significant decrease in the BALF macrophage cell percentages in menthol-flavored e-cig aerosol exposed mice lungs as compared to air controls for both male and female mice (Figure 3F). Of note, though not significant, we saw a decrease in the levels of resident macrophages in the lung digests from mouse lungs exposed to menthol flavored e-cig aerosols with a more pronounced decrease observed amongst female mice (Figure S2A). Interestingly we found a distinct population of RAM expressing markers including Mki67, Pclaf and Top2a. These genes are responsible for cell proliferation. Importantly, the cell scoring for the G2M phase of the cell cycle was high for this cell cluster and thus these were labeled as ‘proliferating’ RAMs (pRAM) (23, 24). The pRAMs showed significant (p=0.0156) changes in the cell percentages with respect to varying exposures in our study (Figures 3D, Table S2).

Removed in the revised version: Taken together, we show a significant shift towards the neutrophil-mediated immune response in the lungs of animals exposed to tobacco-flavored e-cig aerosols in our study. It is important to mention that due to the presence of nicotine in all the e-liquids used as treatment groups for this study, we added an extra control group-PGVG+Nic-for select experiments. Flow cytometric analyses revealed slight variations in the lung neutrophils and macrophage percentages observed in PGVG+Nic group as compared to PGVG only. However, none of these changes were significant. Furthermore, the patterns of change observed for the lungs exposed to aerosols from flavored e-liquids were quite distinct from those observed for PGVG+Nic thus proving that the observed changes are not solely due to the presence of nicotine in these treatment groups (Figures 3E-F, S1A).

Removed in the revised version: Differential gene expression analyses on the myeloid cell cluster identified maximum gene dysregulation in myeloid cell cluster from mouse lungs exposed to tobacco-flavored e-cig aerosols. Interestingly, we did not observe major dysregulation in the gene expression for myeloid cell cluster in fruit-flavored e-cig exposed mouse lungs as compared to air controls. A total of 24 genes (17 upregulated; 7 downregulated) were differentially expressed in the myeloid clusters of lungs of animals exposed to fruit-flavored e-cig aerosols (Figure 4A). GO analyses identified terms like ‘myeloid cell differentiation’, ‘gas transport’ and ‘H2O2 catabolic process’ as the top hits for this cluster (Figures 4B). Importantly, few gene clusters (Hbb-bs, Hbb-bt, Hba-a2 and Hbb-a1) showed sex-specific changes in the gene expressions in exposed lungs as compared to control. These genes clusters enrich for ‘erythrocyte development’ upon GO analyses and warrant further study.

Removed in the revised version: Exposure to flavored e-cig aerosols result in dampening of neutrophil-mediated innate immune responses in C57BL/6J mouse lungs.

Removed in the revised version: Heatmap and bar plot showing the DESeq2 and GO analyses results respectively from the DEGs in the myeloid cell cluster from Fruit (A & B), menthol (C&D) and tobacco (E&F)-flavored e-cig aerosol exposed mouse lungs as compared to air controls. Heatmap (G) and bar plot (H) showing the DESeq2 and GO analyses results respectively from the DEGs in the Ly6G-neutrophil cluster from air and flavored e-cig (fruit, menthol and tobacco) aerosol exposed mouse lungs.

Removed in the revised version: We demonstrated an upregulation of 220 genes and a downregulation of 104 genes in the lungs of mice exposed to menthol-flavored e-cig aerosol as compared to ambient air. We observed increased expression of inflammatory genes including Cxcl3, Il1r2, Stat4, pointing towards the activation of chemokine-mediated signaling due to exposure to menthol-flavored e-cig aerosol. We also found a significant decrease in the expression of Edn1 (endothelin1) in the myeloid cells of flavored e-cig aerosols as compared to air controls (Figure 4C). These genes play a role in positive regulation of neutrophil recruitment and inflammatory responses into the lungs (25–29). GO analyses of the DEGs (p<0.01), thus, showed ‘cytokine-mediated signaling’, ‘cell chemotaxes’, and ‘negative regulation of MAPK cascade’ as the top hits as shown in the Figure 4D.

Removed in the revised version: Like menthol-, tobacco-flavored e-cig aerosol also elicited a significant increase in the expression of 338 genes and a decrease in 215 genes as compared to air controls in the myeloid cell cluster. We observed an increase in the expression of chemokines like Stat4, Il1b, and Il1bos in the myeloid cells resulting in terms like ‘cytokine-mediated signaling pathway’ enriched upon gene enrichment analyses. Like menthol-exposed mouse lungs, we found significant downregulation of Edn1 gene in the lungs of mice exposed to tobacco-flavored e-cig aerosol (Figure 4E). We provide the list of all DEGs as obtained from DESeq2 analyses and GO analyses results in Supplementary Excel Files 1-2.

Removed in the revised version: Based on previous literature (30, 31) and enrichment of genes responsible for neutrophil chemotaxes in the current work, we reckoned this cell type to be of importance to be looked at separately in vivo. As mentioned earlier, we identified two distinct phenotypes of neutrophils-(a) Ly6G+ neutrophils and (b) Ly6G− neutrophils. Interestingly, the percentage increase in the Ly6G− neutrophil population of aerosol exposure was higher as compared to that for Ly6G+ neutrophils per scRNA cell frequency analyses.

Removed in the revised version: Upon performing DESeq2, we did not find any changes in the gene expressions for any of the two neutrophil populations for mouse lungs exposed to menthol-flavored e-cig aerosols despite an increase in the cell percentages of Ly6G-neutrophil phenotype. Interestingly, exposure to fruit-flavored e-cig aerosol caused most changes in the gene expressions in neutrophil population, with Ly6G− neutrophils being affected more than the Ly6G+ population. We observed a dysregulation of 33 and 25 genes in Ly6G− neutrophil cluster from fruit- and PG:VG-aerosol exposed mouse lungs respectively. The Egl-9 Family Hypoxia Inducible Factor 3 (Egln3) gene was found to be downregulated upon exposure to e-cig aerosols in a flavor-independent manner. Egln3 has been known to play a crucial role in lung cancer and its inactivation exhibits lower tumor burden by induction of cancer senescence (32). We further identified dysregulation in the expressions of genes including Irak2, Egr1, Il1r2, Pik3ap1, and Hspa8 in Ly6G− neutrophils for both fruit- and tobacco-aerosol exposed mouse lungs as compared to air controls. Furthermore, upregulation of genes like Irak2 and Egr1 in the Ly6G− neutrophils from mouse lungs exposed to PG:VG and fruit-flavored e-cig aerosol points towards the involvement of NF-κB mediated signaling in our study (Figure 4G). On performing gene enrichment analysis on the dysregulated gene sets from Ly6G-neutrophil cluster we obtained hits including ‘interleukin-1 mediated signaling pathway’, ‘hydrogen peroxide catabolic process’, and ‘pattern recognition signaling pathway’ (Supplementary Excel File 3). Of note, most of the associated genes were upregulated in lungs exposed to fruit-flavored e-cig aerosols. However, the expression pattern of these gene sets points towards dampening of neutrophil-mediated immune responses for lungs exposed to tobacco-flavored e-cig aerosols as compared to air controls (Figure 4H). Compared to Ly6G− neutrophils, we did not observe any change in the gene expression profile of Ly6G+ neutrophils for mouse lungs exposed to tobacco-flavored e-cig aerosol. For both, PG:VG and fruit-flavor e-cig aerosol exposed mouse lungs, Treml4 (Triggering Receptor expressed on Myeloid Cells-Like Protein 4) gene was found to be downregulated in the Ly6G+ neutrophils (Figure S2B). This gene functions as a positive regulator of TLR7 signaling (33, 34).

Activation of T-cell cytotoxic responses in lymphoid cells upon exposure to tobacco-flavored e-cig aerosols

Earlier version: We identified 10 clusters upon sub clustering the lymphoid cell population. These comprised of: Pro-B cell, Pre-B cell, Mature B-cell, CD4+ T-cell, CD8+T-cell, Treg cell, AB T-cell, GD T-cell and Natural killer (NK) cells (35–37). Based on the expression of, Ifit3, Oasl2 and Rsad2 a unique cluster of IFN-stimulated B-cells was identified. The cell frequency of these cells was markedly lowered (19%) in mouse lungs exposed to tobacco-flavored e-cig aerosols as compared to air controls. Two-way ANOVA analysis of scRNA seq data showed significant changes (p = 0.0237) in the cell frequencies of CD4 T cells in a flavor-dependent manner in our treatment groups (Table S2). We also identified sex-dependent changes in the cell frequencies of B-cells (p= 0.0148), pro-B cells (p= 0.0078) and CD8 T cells (p= 0.0265) of the lymphoid cluster in our study (Figures 5A-C, Table S2). Flow cytometric analyses showed a significant increase in the CD4+ T-cell percentages in the lung digest from female mice exposed to tobacco-flavored e-cig aerosol. We also observed a significant increase in the CD8+ T cell percentages in the lungs of tobacco-flavored e-cig aerosol exposed mice lungs as compared to air control in both male and female mice (Figure 5D). Of note, however, we found significant sex-dependent associations between the CD4+ T cell percentages (p=0.0339) upon two-way ANOVA analyses of the flow cytometry data.

Revised version: We next studied the changes in the lymphoid clusters of treated and control samples. Both scRNA seq and flow cytometric analyses showed no variation in the CD4+ T cell frequency in the lungs of different e-cig flavored exposed mouse lungs as compared to air controls. We also identified flavor-dependent variations in the cell frequencies of CD8 T cells (p= 0.0802) in the lymphoid cluster from mouse lungs (Figures 5A-B, Supplementary File S3B). Flow cytometric analyses corroborated with the scRNA seq findings and showed a significant increase in the CD8+ T cell percentages in the lungs of tobacco-flavored e-cig aerosol exposed mice lungs as compared to air control in male mice (Figure 5C-D, Supplementary File S3D). Of note, we found a significant increase in the CD4+ T cell percentages in menthol flavored e-cig aerosol exposed mouse lung as compared to PG:VG exposed control (Figure S5).

Added in the revised version: scRNA seq analyses and flow cytometry result show flavor-dependent increase in CD8 T cells in lungs of differently flavored e-cig aerosol exposed C57BL/6J mouse.

Removed in the revised version: Exposure to tobacco-flavored e-cig aerosols results in significant increase in lung T-cell percentages in C57Bl/6J mice.

Earlier version: A total of 10 unique cell clusters were identified from the lymphoid cell subset of air and e-cig aerosol exposed mouse lung (A). The heatmap showing the most DEGs in each cluster across the myeloid subset. These differential expressions were used to determine and name the cell types in this cluster. The intensity of expression is indicated by the purple coloring (B). The cell percentages of different cell types in the lymphoid cluster showing the sex-dependent variations in the cell composition on exposure to differently flavored e-cig aerosols (C). Flow cytometry results showing changes in the percentages of lung CD4+ and CD8+ T-cells in the lung tissue digest (D) from mice exposed to differently flavored e-cig aerosols as depicted in a sex-specific manner. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, **p<0.01, and ****p<0.0001, per Tukey post hoc two-way ANOVA comparison.

Revised version: The cell frequencies (denoted as counts) of CD4 (A) and CD8 (B) T-cells showing the sex-dependent variations in the cellular composition upon exposure to differently flavored e-cig aerosols. Representative flow plots (C) and bar graph (D) showing changes in the mean counts of CD4+ and CD8+ T-cells in the lung tissue digest from male and female mice exposed to differently flavored e-cig aerosols as determined using flow cytometry. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, and **p<0.01, per Tukey post hoc two-way ANOVA comparison.

Added in the revised version: Exposure to fruit-flavored e-cig aerosol affected the mitotic pathway genes in the lymphoid cluster

Earlier version: We found dysregulation of 112 genes (40 up- and 72-down-regulated) in mouse lungs exposed to fruit-flavored e-cig aerosols in the lymphoid cluster. We observe dysregulation of genes (Incenp, Wnt4, Ccnb2, Trat1, Themis) enriched for ‘T-cell receptor signaling’ and ‘nuclear division’, in exposed group as compared to air control (Figure 6A-B). Upregulation of genes like Malt1 (Mucosa-associated lymphoid tissue lymphoma translocation protein 1), Ppp1r2 (Protein phosphatase inhibitor 2), and Spib (Transcription factor Spi-B) indicates activation of T-cell mediated immune response upon exposure to fruit-flavored aerosol (38–40).

Revised version: To assess the effects of acute exposures to differently flavored e-cig aerosols on the cell composition in the mouse lungs, we performed differential expression analyses for each flavor in comparison to air and PG:VG controls. The mango-flavored e-cig aerosol exposure had the mildest effect on the cell compositions and gene expression as compared to the controls in our study. We did not observe major dysregulation in the gene expression for myeloid cell cluster in fruit-flavored e-cig exposed mouse lungs as compared to air controls. A total of 24 genes (17 upregulated; 7 downregulated) were differentially expressed in the myeloid clusters in lungs of animals exposed to fruit-flavored e-cig aerosols (Figure 6Ai, Supplementary File S5A). GO analyses identified terms like ‘myeloid cell differentiation’, gas transport’ and ‘H2O2 catabolic process’ as the top hits for this cluster (Figures 6Aii, Supplementary File 5B). Importantly, few gene clusters (Hbb-bs, Hbb-bt, Hba-a2 and Hbb-a1) showed sex-specific changes in the gene expressions in exposed lungs as compared to control. These gene clusters enrich for ‘erythrocyte development’ upon GO analyses and warrant further study. We further found dysregulation of 112 genes (40 up- and 72-down-regulated) in mouse lungs exposed to fruit-flavored e-cig aerosols in the lymphoid cluster. Dysregulation of genes (Incenp, Wnt4, Ccnb2, Trat1, Themis) enriched for ‘T-cell receptor signaling’ and ‘nuclear division’ were observed in exposed group as compared to air control (Figure 6Bi, Supplementary file 9A). Upregulation of genes like Malt1 (Mucosa-associated lymphoid tissue lymphoma translocation protein 1), Ppp1r2 (Protein phosphatase inhibitor 2), and Spib (Transcription factor Spi-B) indicates activation of T-cell mediated immune response upon exposure to fruit-flavored aerosol (24–26). GO analyses showed terms like ‘mitotic nuclear division’, ‘mitotic spindle assembly’, and ‘nuclear chromosome segregation’ to be highly affected on exposure to fruit-flavored e-cig aerosol in C57BL/6J mouse lungs (Figure 6Bii, Supplementary File 9B).

Earlier version: Exposure to tobacco flavored e-cig aerosols result in activation of T-cytotoxic responses in C57BL/6J mouse lungs.

Revised version: Exposure to fruit flavored e-cig aerosols result in activation of oxidative stress-mediated innate immunity in C57BL/6J mouse lungs.

Earlier version: Heatmap and bar plot showing the DESeq2 and GO analyses results respectively from the DEGs in the lymphoid cell cluster from Fruit (A & B), menthol (C&D) and tobacco (E&F)-flavored e-cig aerosol exposed mouse lungs as compared to air controls.

Revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to fruit-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (fruit-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from fruit-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Added in the revised version: Exposure to Menthol-flavored e-cig aerosols affected the immune cell function

Added in the revised version: We demonstrated an upregulation of 220 genes and a downregulation of 104 genes in the myeloid cluster of mouse lungs exposed to menthol-flavored e-cig aerosol as compared to ambient air. We observed increased expression of inflammatory genes including Cxcl3, Il1r2, Stat4, pointing towards the activation of chemokine-mediated signaling due to exposure to menthol-flavored e-cig aerosol. We also found a significant decrease in the expression of Edn1 (endothelin1) in the myeloid cells of flavored e-cig aerosols as compared to air controls. These genes play a role in positive regulation of neutrophil recruitment and inflammatory responses into the lungs (27–31). GO analyses of the DEGs (p<0.01), thus, showed ‘cytokine-mediated signaling’, ‘cell chemotaxes’, and ‘negative regulation of MAPK cascade’ as the top hits as shown in Figure 7A (Supplementary File 6A-B).

Added in the revised version: Exposure to Menthol flavored e-cig aerosols result in activation of innate immune responses C57BL/6J mouse lungs.

Added in the revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to menthol-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (menthol-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from menthol-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Earlier version: Contrary to the responses observed for exposure to fruit-flavored e-cig aerosols, we found significant upregulation in the expression of Cdk8 and Camk1d genes in the lymphoid cell populations for menthol-flavored aerosol exposed mouse lungs (Figure 6C-D). Cdk8 (cyclin-dependent kinase 8) is a transcriptional regulator that has a role in the cell cycle progression (41). Whereas Camk1d (calcium/calmodulin dependent protein kinase ID) functions to regulate calcium-mediated granulocyte function and respiratory burst within the cells(42). Taken together, our results point towards an increase in cell proliferation and gene transcription in the lymphoid cluster of mouse lungs exposed to menthol-flavored e-cig aerosols as indicated by enrichment of terms like ‘cyclin-dependent protein serine/threonine kinase activity’, ‘RNA polymerase II CTD modifying activity’ and ‘calmodulin-dependent protein kinase activity’ associated with this cell cluster.

Revised version: Contrary to the responses observed for exposure to fruit-flavored e-cig aerosols, we found significant upregulation in the expression of Cdk8 and Camk1d genes in the lymphoid cell populations for menthol-flavored aerosol exposed mouse lungs (Figure 7B, Supplementary File 10A). Cdk8 (cyclin-dependent kinase 8) is a transcriptional regulator that has a role in the cell cycle progression (32). Whereas Camk1d (calcium/calmodulin dependent protein kinase ID) functions to regulate calcium-mediated granulocyte function and respiratory burst within the cells(33). Taken together, our results point towards an increase in cell proliferation and gene transcription in the lymphoid cluster of mouse lungs exposed to menthol-flavored e-cig aerosols as indicated by enrichment of terms like ‘cyclin-dependent protein serine/threonine kinase activity’, RNA polymerase II CTD modifying activity’ and ‘calmodulin-dependent protein kinase activity’ associated with this cell cluster (Figure 7Bii, Supplementary File 10B).

Added in the revised version: Exposure to Tobacco-flavored e-cig aerosol elicits immune response in myeloid cell and cell cycle arrest in lymphoid cell population

Added in the revised version: Like menthol-, tobacco-flavored e-cig aerosol also elicited a significant increase in the expression of 338 genes and a decrease in 215 genes as compared to air controls in the myeloid cell cluster. We observed an increase in the expression of chemokines like Stat4, Il1b, and Il1bos in the myeloid cells resulting in terms like ‘cytokine-mediated signaling pathway’ enriched upon functional annotation of DEGs. Like menthol-exposed mouse lungs, we found significant downregulation of Edn1 gene in the myeloid cluster in lungs of mice exposed to tobacco-flavored e-cig aerosol (Figure 8A, Supplementary File 7A-B).

Added in the revised version: Exposure to Tobacco flavored e-cig aerosols result in activation of cytolysis and neutrophil chemotaxis in C57BL/6J mouse lungs.

Added in the revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to tobacco-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (tobacco-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from tobacco-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Earlier version: We observed the downregulation of genes responsible for chaperone-mediated protein folding (Cct5, Cct7, Cct8) in the lymphoid cells from tobacco-flavored e-cig aerosol exposed mouse lungs. Downregulation of these genes could be indicative of the accumulation of misfolded proteins in these lungs which may lead to enhanced cell death (43, 44). In fact, we found increased expression of killer cell lectin-like receptor (Klra)-4 and 8 in exposed mice lungs as compared to air controls, thus indicating that upon exposure to tobacco-flavored e-cig aerosols, the lymphoid cells undergo protein misfolding thereby resulting in increased cell death (Figure 6E-F). A detailed account of the DESeq2 results and GO analyses for all the groups is provided in Supplementary Excel files 4-5.

Revised version: We also observed a downregulation of genes responsible for chaperone-mediated protein folding (Cct5, Cct7, Cct8) in the lymphoid cells from tobacco-flavored e-cig aerosol exposed mouse lungs. Downregulation of these genes could be indicative of the accumulation of misfolded proteins in these lungs which may lead to enhanced cell death (34, 35). In fact, we found increased expression of killer cell lectin-like receptor (Klra)-4 and 8 in exposed mice lungs as compared to air controls, thus indicating that upon exposure to tobacco-flavored e-cig aerosols, the lymphoid cells undergo protein misfolding thereby resulting in increased cell death (Figure 8B, Supplementary File 10A-B).

Dysregulation of chemokine signaling and T-cell activation on exposure to flavored e-cig aerosols

Earlier version: Since we showed increased production of cytokines/chemokines, driving the immune responses in mouse lungs exposed to flavored e-cigs, we performed multianalyte assay to determine the levels of these inflammatory cytokines in the lung digests from the exposed animals as shown in Figure 7A. Exposure to tobacco-flavored e-cig aerosol resulted in a marked increase in the levels of chemotactic chemokines including CXCL16, CXCL12, CXC3R and proinflammatory cytokines including CCl12, CCL17, CCL24, and Eotaxin in the mouse lung digests as compared to air control. Supporting our transcriptional data, the levels of IL1b were dampened in the exposed mouse lungs thus showing negative regulation of IL-1 mediated signaling. We identified sex-dependent changes in the cytokine levels in lung digests from fruit-flavored e-cig aerosols with male mice showing more dampening of cytokine/chemokine protein levels as compared to their female counterparts. Interestingly, the fold changes in the PGVG+Nic group were contrasting to those observed by PGVG alone and flavored e-cig aerosol exposed mouse lungs, again supporting our previous deduction that the observed changes are observed as a concerted effect of chemical present in the e-liquid used for flavored e-cig aerosol exposure.

Revised version: Since we showed increased production of cytokines/chemokines, driving the immune responses in mouse lungs exposed to flavored e-cigs, we performed multianalyte assay to determine the levels of these inflammatory cytokines in the lung digests from the exposed animals as shown in Figure S3D. Exposure to tobacco-flavored e-cig aerosol resulted in a marked increase in the levels of chemotactic chemokines including CXCL16, CXCL12, CXC3R and proinflammatory cytokines including CCl12, CCL17, CCL24, and Eotaxin in the mouse lung digests as compared to air control. Supporting our transcriptional data, the levels of IL1b were dampened in the exposed mouse lungs thus showing negative regulation of IL-1 mediated signaling. We identified sex-dependent changes in the cytokine levels in lung digests from fruit-flavored e-cig aerosols with male mice showing more dampening of cytokine/chemokine protein levels as compared to their female counterparts which aligns with lowed innate immune responses driven by neutrophils and macrophages in the males. Interestingly, the fold changes in the PG:VG+Nic group were contrasting to those observed by PG:VG alone and flavored e-cig aerosol exposed mouse lungs, again supporting our previous deduction that the observed changes are observed as a concerted effect of chemical present in the e-liquid used for flavored e-cig aerosol exposure.

Removed in the revised version: GO analyses of the commonly dysregulated genes show enrichment of dampened innate immune response upon exposure to flavored e-cig aerosols in C57BL/6J mice.

Removed in the revised version: Heat map showing the sex-dependent average fold changes in the levels of cytokines/chemokines in the lung of differently flavored e-cig aerosol exposed C57Bl/6J mouse lungs as compared to air controls. Data are shown as mean ± SEM (n = 3/sex/group) (A). Heatmap showing the fold changes in the expression of commonly dysregulated genes in the myeloid cluster and lymphoid clusters in mouse lungs exposed to flavored e-cig aerosols as compared to ambient air as determined after DESeq2 analyses (B). CNET plot results showing the pathways regulated by the common DEGs on acute (5-day) exposure to differently flavored fruit, menthol, and tobacco) e-cig aerosols in C57Bl/6J mouse lungs (B).

Earlier version: To identify genes that were commonly dysregulated upon exposure, we generated a list of common genes that were significantly dysregulated in exposure categories (fruit, menthol and tobacco). We identified 9 such target genes-Neurl3, Egfem1, Stap1, Tfec, Mitf, Cirbp, Hist1h1c, Gmds, and Htr2c that were dysregulated in the myeloid cluster from lungs exposed to differently flavored e-cig aerosol, but not PGVG. We observed significant upregulation of Neurl3, Stap1, Cirbp, and Hist1h1c and downregulation of Tfec, Mitf, Gmds, and Htr2c in the myeloid cluster of mice exposed to differently flavored e-cig aerosol as compared to air control (Figure 7B). On analyzing the lymphoid cluster for commonly dysregulated genes, we identified - Klra8 (Killer cell lectin-like receptor 8) and Nfia (nuclear factor I)-that were significantly upregulated in the exposure groups as compared to air-controls (Figure 7E). Klra8 is a natural killer cell associated gene, and its upregulation is generally associated with viral infection associated host immune responses within the mouse lungs (45–47). Nfia is a transcriptional activator responsible for regulating Oxphos-mediated mitochondrial responses and proinflammatory pathways (48, 49).

Revised version: To identify genes that were commonly dysregulated upon exposure, we generated a list of common genes that were significantly dysregulated in exposure categories (fruit, menthol and tobacco). We identified 9 such target genes-Neurl3, Egfem1, Stap1, Tfec, Mitf, Cirbp, Hist1h1c, Gmds, and Htr2c that were dysregulated in the myeloid cluster from lungs exposed to differently flavored e-cig aerosol, but not PG:VG. We observed significant upregulation of Neurl3, Stap1, Cirbp, and Hist1h1c and downregulation of Tfec, Mitf, Gmds, and Htr2c in the myeloid cluster of mice exposed to differently flavored e-cig aerosols as compared to air control (Supplementary Figure 6A). On analyzing the lymphoid cluster for commonly dysregulated genes, we identified - Klra8 (Killer cell lectin-like receptor 8) and Nfia (nuclear factor I)- that were significantly upregulated in the exposure groups as compared to air-controls (Figure S6A). Klra8 is a natural killer cell associated gene, and its upregulation is generally associated with viral infection associated host immune responses within the mouse lungs (36–38). Nfia is a transcriptional activator responsible for regulating Oxphos-mediated mitochondrial responses and proinflammatory pathways (39, 40).

Earlier version: Overall, we identified a total of 29 commonly dysregulated gene targets were identified from five major cell clusters and performed gene enrichment analyses on the identified targets to identify the top hits (Tables 4-5). Terms like ‘negative regulation of immune system’ (Hmgb3/Gpam/Scgb1a1/Stap1/Ldlr), ‘positive regulation of lipid biosynthetic pathway’ (Htr2c/Gpam/Ldlr), and ‘receptor recycling’ (Ldlr/Ramp3) were amongst the top hits in our observations (Figure 7C).

Revised version: Overall, we identified a total of 29 commonly dysregulated gene targets that were identified from five major cell clusters and performed gene enrichment analyses on the identified targets to identify the top hits (Tables 2-3). Terms like ‘negative regulation of immune system’ (Hmgb3/Gpam/Scgb1a1/Stap1/Ldlr), ‘positive regulation of lipid biosynthetic pathway’ (Htr2c/Gpam/Ldlr), and ‘receptor recycling’ (Ldlr/Ramp3) were amongst the top hits in our observations (Figure S6B).

Earlier version: Of note, the data presented in this study is a sub-part of a larger study. In addition to the groups mentioned in this manuscript, we also had two additional groups of Tobacco-Derived (TDN) and Tobacco-Free Nicotine (TFN). Though further objectives and experimentations performed in both these studies were varied, common air and PG:VG samples were used for analyses of cytokine/chemokine and cell count data as described in our recent publication(50).

Revised version: Of note, the data presented in this study is a sub-part of a larger study. In addition to the groups mentioned in this manuscript, we also had two additional groups of Tobacco-Derived (TDN) and Tobacco-Free Nicotine (TFN). Though further objectives and experimentations performed in both these studies were varied, common air and PG:VG samples were used for analyses of cytokine/chemokine and cell count data as described in our recent publication(41).

Discussion

Earlier version: E-cigs and associated products have constantly been under scrutiny by the US Food and Drug Administration (FDA) due to public health concerns. In February 2020, FDA placed a regulation on all cartridge-based flavored e-cigs except for menthol and tobacco to reduce the use of e-cigs amongst adolescents and young adults. But it left a loophole for the sale of flavored (including menthol) disposable and open system e-cigs (1, 51). Importantly, most e-cig related bans in the US have happen at the state level, thus allowing differential levels of restrictions imposed on the premarket tobacco applications (PMTAs) and sales which defeats the purpose of limiting their accessibility to the general public (52, 53). Each year new products are introduced in the market with newer device designs and properties, to lure the users (adults between the ages of 18-24 years) which makes it crucial to continue with the assessments of toxicity and health effects of e-cigs in an unbiased manner (54).

Revised version: E-cigs and associated products have constantly been under scrutiny by the US Food and Drug Administration (FDA) due to public health concerns. In February 2020, FDA placed a regulation on all cartridge-based flavored e-cigs except for menthol and tobacco to reduce the use of e-cigs amongst adolescents and young adults. But it left a loophole for the sale of flavored (including menthol) disposable and open system e-cigs (1, 42). Importantly, most e-cig related bans in the US happened at the state level, thus allowing differential levels of restrictions imposed on the premarket tobacco applications (PMTAs) and sales which defeats the purpose of limiting their accessibility to the general public (43, 44). In fact, the use of nicotine containing e-cigs amongst youth and associated policy restrictions have recently been found to be linked to unintended increase in traditional cigarette use (45). Each year new products are introduced in the market with newer device designs and properties, to lure the users (adults between the ages of 18-24 years) which makes it crucial to continue with the assessments of toxicity and health effects of e-cigs in an unbiased manner (46).

Earlier version: Numerous studies indicate increased oxidative stress, DNA damage, and loss of neutrophil function due to exposure to e-cig aerosols in vitro and in vivo (7, 30, 31, 50, 55, 56). However, we do not have much knowledge about the cell populations and biological signaling mechanisms that are most affected upon exposure to differently flavored e-cig aerosols. To bridge this gap in knowledge, we studied the transcriptional changes in the inflammatory responses due to acute (1-hr nose-only exposure including 120 puffs for 5-consecutive days) exposure to fruit-, menthol- and tobacco-flavored e-cig aerosols using single-cell technology.

Revised version: Numerous studies indicate increased oxidative stress, DNA damage, and loss of neutrophil function due to exposure to e-cig aerosols in vitro and in vivo (7, 41, 47–50). However, we do not have much knowledge about the cell populations and biological signaling mechanisms that are most affected upon exposure to differently flavored e-cig aerosols at a single-cell level. To bridge this gap in knowledge, we studied the transcriptional changes in the inflammatory responses due to acute (1-hr nose-only exposure including 120 puffs for 5-consecutive days) exposure to fruit-, menthol- and tobacco-flavored e-cig aerosols using single-cell technology.

Earlier version: Reports indicate that the release of metal ions due to the burning of metal coil is a major source of variation during e-cig exposures (57–59). A 2021 study reported the presence of 21 elements in the pod atomizers from different manufacturers identifying a high abundance of 11 elements including nickel, iron, zinc, and sodium amongst others (60). Previous studies have also shown the presence of similar elements in the e-cig aerosols which could have a possible adverse health effect on the vapers (15, 61). But more importantly, our study points towards a much graver issue, which pertains to the product design of the e-cig vapes. We believe that the aerosol composition varies based on the type of atomizer, coil resistance, coil composition, and chemical reactivity of the e-liquid being used. While much work is done on the chemical composition of the flavors and e-liquids, the other aspects of device design remain understudied and must be an area of research in the future. Another factor that may limit the interpretation of our results pertains to the correlation between the leached metal and the observed transcriptional changes. Since our study provides proof of day-to-day variation in the leaching of metal ions from the same liquid using the same atomizer, it could be possible to develop a statistical model correlating the differential metal exposure to the gene expression changes. We do not conduct such analyses as this was not the focus of this work, but it is a possibility that could be explored in the future.

Revised version: Reports indicate that the release of metal ions due to the burning of metal coil is a major source of variation during e-cig exposures (51–53). A 2021 study reported the presence of 21 elements in the pod atomizers from different manufacturers identifying a high abundance of 11 elements including nickel, iron, zinc, and sodium amongst others (54). Previous studies have also shown the presence of similar elements in the e-cig aerosols which could have a possible adverse health effect on the vapers (19, 55). More importantly, our study points towards a much important issue, which pertains to the product design of the e-cig vapes. We believe that the aerosol composition varies based on the type of atomizer, coil resistance, coil composition, and chemical reactivity of the e-liquid being used. While much work is done on the chemical composition of the flavors and e-liquids, the other aspects of device design remain understudied and must be an area of research in the future. Another factor that may limit the interpretation of our results pertains to the correlation between the leached metal and the observed transcriptional changes. Since our study provides proof of day-to-day variation in the leaching of metal ions from the same liquid using the same atomizer, it could be possible to develop a statistical model correlating the differential metal exposure to the gene expression changes. We do not conduct such analyses as this was not the focus of this work, but it is a possibility that could be explored in the future.

Earlier version: E-cig vaping has been known to affect the innate and adaptive immune responses among vapers (30, 62–64), but flavor-specific effects on immune function are not fully explored. While we expected to see flavor-dependent changes in our experiment, we did not anticipate observing interesting sex-dependent variation in the lung tissues at single-cell level. In this respect, recent studies show concurring evidence suggesting sex-specific changes in lung inflammation, mitochondrial damage, gene expression, and even DNA methylation in mice exposed to e-cig aerosols (65, 66).

Revised version: E-cig vaping has been known to affect the innate and adaptive immune responses among vapers (49, 56–58), but flavor-specific effects on immune function are not fully explored. While we expected to see flavor-dependent changes in our experiment, we did not anticipate observing interesting sex-dependent variation in the lung tissues at single-cell level. In this respect, recent studies show concurring evidence suggesting sex-specific changes in lung inflammation, mitochondrial damage, gene expression, and even DNA methylation in mice exposed to e-cig aerosols (59, 60).

Earlier version: One of the most interesting discoveries from our single-cell analyses was the identification of a new cluster of neutrophils that was called Ly6G-neutrophils in this study. While we observed an increase in the cell percentages of these cells in our treatment groups, little to no change in the gene expression was noted in this study which could be indicative of impaired function of these neutrophil population, a fate being reported by various previous studies pertaining to e-cig exposures (13, 62). In fact, our study identified this unique population and showed that IL-1 mediated responses were dampened in tobacco and menthol-exposed lung neutrophils which was corroborated by mild to no change in the levels of IL-1β in the lung tissue digests from exposed mice as compared to air. Ly6G is an important marker of neutrophil maturation in mammalian cells and has been reported in relation to various bacterial and parasitic infections in previous studies (67, 68). Exposure to e-cig aerosols may result in a loss of function of mature neutrophils in the mouse lungs which gets actively replaced by inactive/immature neutrophils that are deficient in Ly6G expression. In fact, this could explain our inability to identify significant changes in mature neutrophil population through flow cytometry (where Ly6G was used as a marker) when comparing treated and control groups. Interestingly, our in vivo findings were corroborated by a recently published work where neutrophils from healthy volunteers demonstrated a reduction in neutrophil chemotaxis, phagocytic function, and neutrophil extracellular trap formation on exposure to e-cig aerosols, thus suggesting loss of function by mature neutrophils (30). However, this study did not investigate the effects on the immature neutrophil population, an area that needs further study in future research.

Revised version: One of the most interesting discoveries from our single-cell analyses was the identification of a cluster of immature neutrophils without Ly6G surface marker. While we observed an increase in the cell percentages of these cells in our treatment groups, little to no change in the gene expression was noted which could be indicative of impaired function of these neutrophil population, a fate being reported by various previous studies pertaining to e-cig exposures (13, 56). In fact, our study identified an increase in the level of both mature (Ly6G+) and immature (Ly6G-) neutrophils in tobacco-exposed mouse lung which was corroborated by (a) increase in staining for Ly6G – and Ly6G+ neutrophils in female, but not male mice, through immunofluorescence, and (b) mild to no change in the levels of IL-1β in the lung tissue digests from exposed mice as compared to air. Ly6G is an important marker of neutrophil maturation in mammalian cells and has been reported in relation to various bacterial and parasitic infections in previous studies (61, 62).Exposure to tobacco flavored e-cig aerosols may result in a loss of function of mature neutrophils in the mouse lungs which gets replaced by inactive/immature neutrophils that are deficient in Ly6G surface marker for maturation. In fact, this could explain our inability to identify significant changes in mature neutrophil population through flow cytometry (where Ly6G was used as a marker) when comparing treated and control groups. Interestingly, our in vivo findings were corroborated by a recently published work where neutrophils from healthy volunteers demonstrated a reduction in neutrophil chemotaxis, phagocytic function, and neutrophil extracellular trap formation on exposure to e-cig aerosols, thus suggesting loss of function by mature neutrophils (49). However, this study did not investigate the effects on the immature neutrophil population, an area that needs further study in future research.

Earlier version: Of note, myeloid and lymphoid limbs of immunity are interconnected (69, 70). A decline in the neutrophil-mediated immunity might activate other cell types to offer protection. In our study, we find a decline in the neutrophilic immune responses in menthol and tobacco flavored e-cig exposed mouse lungs. Contrarily, we report an increase in the T-cell responses in the form of increased CD4+ and CD8+ T cells from both scRNA seq and flow cytometric analyses. In fact, increased expression of genes including Malt1, Serpinb9b, and Sema4c are indicative of enhanced T-cell mediated immune response in the lungs of mice exposed to fruit(mango) flavored e-cig aerosols (71–73). Contrary to this, exposure to menthol-flavored e-cig aerosols had a much milder effect on the lymphoid population within the lungs of C57Bl/6J mice. We found evidence for increased cell proliferation due to activation of cyclin-dependent protein kinase signaling mediated via expression of genes including Cdk8 and Camk1d in these cells (42, 74). Exposure to tobacco-flavored e-cig aerosol provided evidence for decreased chaperone-mediated protein folding, due to the downregulation of Chaperonin Containing TCP-1 (CCT) family of proteins. Chaperonin Containing TCP-1 proteins are important to regulate the production of native active, tubulin, and other proteins crucial for cell cycle progression and cytoskeletal organization (75). This is in conjunction with the upregulation of Klra4 and Klra8 that is indicative of increased protein misfolding and cytotoxic responses in the lymphoid cells of tobacco-exposed e-cig aerosols (76, 77).

Revised version: Of note, myeloid and lymphoid limbs of immunity are interconnected (63, 64). A decline in the neutrophil-mediated immunity might activate other cell types to offer protection. In our study, we find a decline in the neutrophilic immune responses in menthol and tobacco flavored e-cig exposed mouse lungs. Contrarily, we report an increase in the T-cell responses in the form of increased CD8+ T cells from both scRNA seq and flow cytometric analyses in male mice. In fact, increased expression of genes including Malt1, Serpinb9b, and Sema4c are indicative of enhanced T-cell mediated immune response in the lungs of mice exposed to fruit(mango) flavored e-cig aerosols (65–67). Contrary to this, exposure to menthol-flavored e-cig aerosols had a much milder effect on the lymphoid population within the lungs of C57BL/6J mice. We found evidence for increased lymphoid cell proliferation due to activation of cyclin-dependent protein kinase signaling mediated via expression of genes including Cdk8 and Camk1d in these cells (33, 68). Exposure to tobacco-flavored e-cig aerosol provided evidence for decreased chaperone-mediated protein folding, due to the downregulation of Chaperonin Containing TCP-1 (CCT) family of proteins. Chaperonin Containing TCP-1 proteins are important to regulate the production of native active, tubulin, and other proteins crucial for cell cycle progression and cytoskeletal organization (69). This is in conjunction with the upregulation of Klra4 and Klra8 that is indicative of increased protein misfolding and cytotoxic responses in the lymphoid cells of tobacco-exposed e-cig aerosols (70, 71).

Earlier version: Overall, we provide evidence of subdued innate immune responses due to loss of function of neutrophils and increased T-cell proliferation and cytotoxicity in a flavor-dependent manner upon exposure to e-cig aerosols in this study. An increase in the levels of CCL17, CCL20, CCL22, IL2 and Eotaxin in the lung digests from tobacco-exposed mouse lungs further support this deduction as these cytokines/chemokines are associated with T-cell mediated immune responses (78–85). Importantly CXCL16 attracts T-cells and natural killer cells to activate cell death. It has been reported to be involved in LPS-mediated acute lung injury (ALI), an outcome which has also been linked with e-cig exposures in human (86, 87).

Revised version: Overall, we provide evidence of subdued innate immune responses due to loss of function of neutrophils and increased T-cell proliferation and cytotoxicity in a flavor-dependent manner upon exposure to e-cig aerosols in this study. An increase in the levels of CCL17, CCL20, CCL22, IL2 and Eotaxin in the lung digests from tobacco-exposed mouse lungs further support this deduction as these cytokines/chemokines are associated with T-cell mediated immune responses (72–79). Importantly CXCL16 attracts T-cells and natural killer cells to activate cell death. It is involved in LPS- mediated acute lung injury (ALI), an outcome which has also been linked with e-cig exposures in human (80, 81).

Earlier version: In the end, we compiled a list of commonly dysregulated genes in a flavor independent manner and identified 29 gene targets. Signal-transducing adaptor protein-2 (Stap1) which is commonly upregulated upon exposure to e-cig aerosols is known to regulate T-cell activation and airway inflammation which falls in line with the overall crux of our findings (88). Another gene that was found to be consistently dysregulated in many cell types was Cold-inducible RNA binding protein (Cirbp). Cirbp is a stress response proteins linked with stressors like hypoxia. Its upregulation upon e-cig exposure supports that vaping induces oxidative stress and can have adverse implications on the exposed cell types (89). 5-hydroxytryptamine receptor 2C (Htr2c) and Klra8 are other genes in this category of commonly dysregulated genes that are associated with enhancing inflammation and cell death (76, 77, 90). Overall, we provide a cell-specific atlas of immune responses upon exposure to differently flavored e-cig aerosols.

Revised version: Further, we compiled a list of commonly dysregulated genes in a flavor independent manner and identified 29 gene targets. Signal-transducing adaptor protein-2 (Stap1) which is commonly upregulated upon exposure to e-cig aerosols is known to regulate T-cell activation and airway inflammation which is in agreement with the overall outcome of our findings (82). Another gene that was found to be consistently dysregulated in many cell types was Cold-inducible RNA binding protein (Cirbp). Cirbp is a stress response protein linked with stressors like hypoxia. Its upregulation upon e-cig exposure supports that vaping induces oxidative stress and can have adverse implications on the exposed cell types (83). 5-hydroxytryptamine receptor 2C (Htr2c) and Klra8 are other genes in this category of commonly dysregulated genes that are associated with enhancing inflammation and cell death (70, 71, 84). Overall, we provide a cell-specific atlas of immune responses upon exposure to differently flavored e-cig aerosols.

Earlier version: Considering that scRNA technology has not been commonly used for e-cig research, ours is one of the first studies employing this technique to identify possible changes in the cellular composition and gene expressions. Importantly we use the nose-only exposure system for our experiment to avoid exposure through other routes. However, despite the novel approach and state-of-the-art exposure system, we had a few limitations. First, we used a small sample size to identify the changes in the mouse lungs upon exposure to e-cig aerosols at a single-cell level. Due to the expensive nature of single-cell sequencing technology and limited information in the literature reporting changes at single-cell level, we chose to design this experiment with small sizes of experimental and validation cohort. But, based on the encouraging findings from this pilot study, future studies could be designed with a larger sample size and longer durations of exposure to identify the acute and chronic effects of vaping. Second, we could not delineate the sex-dependent changes upon exposure from our data. While we found evidence of sex-dependent differences in immune responses on exposure to differently flavored e-cig aerosols we pooled the samples from 2 mice/sex/group to perform scRNA seq analyses. This limits us from analyzing and drawing conclusions with regards to the sex-dependent variations at the level of differential gene expression in our data with confidence. Thus, we chose to use aggregate analyses to make comparisons. However, future experimental designs must study sex as a confounder for studying the effects of e-cig exposure in humans. Third, the inclusion of PGVG+Nic group was streamlined in this study, but future work including this group for scRNA seq analyses might be crucial to delineate the effects of nicotine alone on gene transcription. Fourth, we did not anticipate changes in the metal release on consecutive days of exposure in our study. Our data pointed towards the importance of device design in e-cig exposures. Future studies need to identify the factors that may affect the daily composition of e-cig aerosols and devise a method of better monitoring these possible confounders. However, in this regard our experiment does mimic the real-life well, as such variations due to prolonged storage of e-liquid and differences arising due to vape design must be common amongst human vapers.

Revised version: Considering that scRNA technology has not been commonly used for e-cig research, ours is one of the first studies employing this technique to identify possible changes in the cellular composition and gene expressions. Importantly we use the nose-only exposure system for our experiment to avoid exposure through other routes. However, despite the novel approach and state-of-the-art exposure system, we had a few limitations. First, we used a small sample size to identify the changes in the mouse lungs upon exposure to e-cig aerosols at a single-cell level. Due to the expensive nature of single-cell sequencing technology and limited information in the literature reporting changes at single-cell level, we chose to design this experiment with small sizes of experimental and validation cohort. But, based on the encouraging findings from this study, future studies could be designed with a larger sample size,longer durations of exposure and more targeted approach to identify the acute and chronic effects of vaping in vivo. Second, we could not expand upon the sex-dependent changes observed through our work upon exposure. This was because such an effect was not anticipated when we conceived the idea of a short-term exposure in mice. However, considering the evidence from the cuttenr study, future experimental designs in our lab will be designed taking sex as a crucial confounder for studying the effects of e-cig exposure in translational contexts. Third, the inclusion of PG:VG + Nic group was streamlined in this study, but in future work including this group for scRNA seq analyses might be crucial to delineate the effects of nicotine alone on gene transcription. Fourth, we did not anticipate changes in the metal release on consecutive days of exposure at the start of our study. Later, our data pointed towards the importance of device design in e-cig exposures. Future studies need to identify the factors that may affect the daily composition of e-cig aerosols and devise a method of better monitoring these possible confounders. However, in this regard our experiment does mimic the real-life scenario, as such variations due to prolonged storage of e-liquid and differences arising due to vape design must be common amongst human vapers.

In conclusion, we identified cell-specific changes in the gene expressions upon exposure to e-cig aerosols using single cell technology. We identified a set of top 29 dysregulated genes that could be studied as markers of toxicity/immune dysfunction in e-cig research. Future work with larger sample sizes and sex-distribution is warranted to understand the health impacts of long-term use of these novel products in humans.

Methods

Animals Ethics Statement

All experiments were conducted per the guidelines set by the University Committee on Animal Resources at the University of Rochester Medical Center (URMC). Care was taken to implement an unbiased and robust approach during the experimental design and conduction of each experiment to ensure data reproducibility per the National Institutes of Health standards.

Animals

We ordered 5-week-old pups of male and female C57BL/6J mice from Jackson Laboratory to conduct this experiment. Prior to the start of the experiment, mice were housed at the URMC Vivarium for acclimatization. Thereafter the animals were moved to the mouse Inhalation Facility at URMC for training and exposures.

Earlier version: One week prior to the start of the exposures, mice underwent a five-day nose-only training to adapt themselves to the mesh restraints of the exposure tower. The mouse restraint durations were increased gradually to minimize the animal stress and discomfort. Of note, the mouse sacrifice was performed within 8-12 weeks’ age for each mouse group to ensure that the mouse age corresponds to the age of adolescents (12-17 years) in humans (91, 92).

Revised version: One week prior to the start of the exposures, mice underwent a five-day nose-only training to adapt themselves to the mesh restraints of the exposure tower. The mouse restraint durations were increased gradually to minimize the animal’s stress and discomfort. Of note, the mouse sacrifice was performed within 8-12 weeks’ age for each mouse group to ensure that the mouse age corresponds to the age of adolescents (12- 17 years) in humans (85, 86). Age and sex-matched animals (n = 2/sex/group) used to perform single cell RNA sequencing (scRNA seq) were considered as the ‘experimental cohort’; whereas another group of age and sex-matched (n = 3/sex/group) mice exposed to air and flavored e-cig aerosol served as ‘validation cohort’ for this study.

E-cigarette Device and E-liquid

Earlier version: We utilized an eVic-VTC mini and CUBIS pro atomizer (SCIREQ, Montreal, Canada) with a BF SS316 1.0-ohm coil from Joyetech (Joyetech, Shenzhen, China) for vaping and the inExpose nose-only inhalation system from SCIREQ (SCIREQ, Montreal, Canada) for mouse exposures. Both air and PG:VG exposed mice groups were considered as controls for this experiment. We used commercially available propylene glycol (PG; EC Blend) and vegetable glycerin (VG; EC Blend) in equal volume to prepare a 50:50 solution of PG:VG. For flavored product exposures, mice were exposed to three different e-liquids – a menthol flavor “Menthol-Mint”, a fruit flavor “Mango” and a tobacco flavor “Cuban Blend”. Of note, all the e-liquids were commercially manufactured with 50mg/mL of tobacco derived nicotine (TDN). So, all treatments have nicotine in addition to the flavoring mentioned respectively. Additionally, we used a mixture of PG:VG with 50mg/mL of TDN as a control for limited experiments to study the effect of nicotine alone in our treatment. This group is labeled as PGVG+Nic for the rest of the manuscript.

Revised version: We utilized an eVic-VTC mini and CUBIS pro atomizer (SCIREQ, Montreal, Canada) with a BF SS316 1.0-ohm coil from Joyetech (Joyetech, Shenzhen, China) for vaping and the inExpose nose-only inhalation system from SCIREQ (SCIREQ, Montreal, Canada) for mouse exposures. Both air and PG:VG exposed mice groups were considered as controls for this experiment. We used commercially available propylene glycol (PG; EC Blend) and vegetable glycerin (VG; EC Blend) in equal volumes to prepare a 50:50 solution of PG:VG. For flavored product exposures, mice were exposed to three different e-liquids – a menthol flavor “Menthol-Mint”, a fruit flavor “Mango” and a tobacco flavor “Cuban Blend”. Of note, all the e-liquids were commercially manufactured with 50mg/mL of tobacco derived nicotine (TDN). So, all treatments have nicotine in addition to the flavoring mentioned respectively. Additionally, we used a mixture of PG:VG with 50mg/mL of TDN as a control for limited experiments to study the effect of nicotine alone in our treatment. This group is labeled as PG:VG+Nic for the rest of the manuscript.

E-cigarette Exposure

Earlier version: Scireq Flexiware software with the InExpose Inhalation system was used for controlling the Joyetech eVic-VTC mini device to perform nose-only mouse exposures. For this exposure, we utilized a puffing profile that mimicked the puffing topography of e-cig users in two puffs per minute with a puff volume of 51 mL, puff duration of 3 seconds, and an inter puff interval of 27 seconds with a 2 L/min bias flow between puffs (93). Figures 1A and 1B depict the experimental design and exposure system employed for this study.

Revised version: Scireq Flexiware software with the InExpose Inhalation system was used for controlling the Joyetech eVic-VTC mini device to perform nose-only mouse exposures. For this exposure, we utilized a puffing profile that mimicked the puffing topography of e-cig users in two puffs per minute with a puff volume of 51 mL, puff duration of 3 seconds, and an inter puff interval of 27 seconds with a 2 L/min bias flow between puffs (87). Figures 1A & S1A depict the experimental design and exposure system employed for this study.

Earlier version: Age-matched male and female (n=5 per sex) mice were used for each group, namely, air, PG:VG, Fruit, Menthol, and Tobacco. To ensure rigor and reproducibility in our work, we have used age- and sex-matched control and treated mice in this study. Confounders like environment and stress were minimized by housing all the cages in environmentally controlled conditions and training all the mice (both control and treated) in nose-only chambers. Each group of mice was exposed to the above-mentioned puffing profile for one hour each day (120 puffs) for a total of five consecutive days. Additionally, a group of mice (n=3/sex) was exposed to PGVG+Nic for the same duration using similar exposure profile to serve as control to assess the effect of nicotine on the observed changes using selected experiments. Air-exposed mice were exposed to the same puffing profile for a total of five consecutive days to ambient air. We recorded the temperature, humidity, and CO levels of the aerosols generated at the start, mid, and end of the exposure on each day using the Q-Trak Indoor Air Quality Monitor 7575 (TSI, Shoreview, MN). Total Particulate Matter (TPM) sampling was done from the exhaust tubing of the set-up at the 30 min mark of the exposure and at the inlet connected to the nose-only tower (shown in Figure 1B) immediately after the culmination of the exposure. Gravimetric measurements for TPM were also conducted to confirm relative dosage to each mouse group daily.

Revised version: Age-matched male and female (n=5 per sex) mice were used for each group, namely, air, PG:VG, Fruit, Menthol, and Tobacco. To ensure rigor and reproducibility in our work, we have used age- and sex-matched control and treated mice in this study. Confounders like environment and stress were minimized by housing all the cages in environmentally controlled conditions and training all the mice (both control and treated) in nose-only chambers. Each group of mice was exposed to the above-mentioned puffing profile for one hour each day (120 puffs) for a total of five consecutive days.

Added in the revised version: Additionally, a group of mice (n=3/sex) was exposed to PG:VG + Nic for the same duration using similar exposure profile to serve as control to assess the effect of nicotine on the observed changes using selected experiments. Air-exposed mice were exposed to the same puffing profile for a total of five consecutive days to ambient air. We recorded the temperature, humidity, and CO levels of the aerosols generated at the start, mid, and end of the exposure on each day using the Q-Trak Indoor Air Quality Monitor 7575 (TSI, Shoreview, MN). Total Particulate Matter (TPM) sampling was done from the exhaust tubing of the set-up at the 30 min mark of the exposure and at the inlet connected to the nose-only tower (shown in Figure S1A) immediately after the culmination of the exposure. Gravimetric measurements for TPM were also conducted to confirm relative dosage to each mouse group daily.

Preparation of single-cell suspension

Earlier version: The animals were sacrificed (at 8-10 weeks’ age) immediately after the final exposure. Vascular lung perfusion was performed using 3mL of saline before harvesting the lung lobes for preparation of single-cell suspension. It is important to mention here that of the 5 lung samples/sex/group; 2 sex/group were used for histological assessments and scRNA analyses. Here, the left lung lobe was inflated using low-melting agarose and used for histology, while the rest of the uninflated lung lobes were used for preparing the single-cell suspension. We pooled the lung lobes for each sex per group for preparation of the single-cell suspension as depicted in Figure 1A. The lung lobes were weighed digested using Liberase method as described earlier(94).

Revised version: The animals were sacrificed (at 8-10 weeks’ age) immediately after the final exposure. Vascular lung perfusion was performed using 3 mL of saline before harvesting the lung lobes for preparation of single-cell suspension. It is important to mention here that of the 5 lung samples/sex/group; 2 sex/group were used for histological assessments and scRNA analyses. Here, the left lung lobe was inflated using low-melting agarose and used for histology, while the rest of the uninflated lung lobes were used for preparing the single-cell suspension. We pooled the lung lobes for each sex per group for preparation of the single-cell suspension as depicted in Figure 1A. The lung lobes were weighed digested using Liberase method as described earlier (88). Briefly, lung lobes were weighed and digested using Liberase (Cat# 5401127001; Roche, Basel, Switzerland) enzymatic cocktail with 1% DNase. The tubes were then transferred to the gentleMACS dissociator (Miltenyi, Gaithersburg, MD) and the manufacturer’s protocol for mouse lung digestion was run. The sample tubes were next incubated at 370C for 30 minutes with constant rotation after which the suspension was strained through 70- micron MACS Smart Strainer. Thereafter, the suspension was centrifuged at 500g for 10 minutes at 40C, the supernatant was discarded, and 0.5 mL of RBC lysis buffer was added to the cell pellet to digest RBCs. The suspension was left on ice for 5 min in RBC lysis buffer and then 4 mL of ice-cold PBS with 10% FBS was added to stop the lysis. The suspension was again centrifuged at 500g for 10 minutes at 40C. The cell pellet was suspended in 1 mL PBS with 10% FBS, and cell number and viability were checked using AO/PI staining on a Nexcelom Cellometer Auto2000.

Library Preparation and Single-Cell Sequencing

Earlier version: The prepared single-cell suspension was sent to the Genomics Research Center (GRC) at URMC for library preparation and single-cell sequencing. Library preparation was performed from control and treatment groups using the 10X single-cell sequencing pipeline by 10X Genomics and 10000 cells were captured per sample using the Chromium platform. The prepared library was sequenced on NovaSeq 6000 (Illumina, San Diego, CA) at a mean sequence depth of 30,000 reads per cell. Read alignment was performed to GRCm38 Sequence.

Revised version: The prepared single-cell suspension was sent to the Genomics Research Center (GRC) at URMC for library preparation and single-cell sequencing. Library preparation was performed from control and treatment groups using the 10X single cell sequencing pipeline by 10X Genomics and 10,000 cells were captured per sample using the Chromium platform. The prepared library was sequenced on NovaSeq 6000 (Illumina, San Diego, CA) at a mean sequence depth of 30,000 reads per cell. Read alignment was performed to GRCm38 Sequence.

Data Analyses

Earlier version: We used the standard Seurat v4.2 analyses pipeline to analyze our data (95). In brief the low-quality cells and potential doublets were excluded from the dataset to create the analyses dataset. The residual features due to the presence of RBCs were corrected before integration. “scTransform” function was used for integration of all the datasets after which the standard Seurat pipeline was used for data normalization of integrated data. “FindVariableGenes” gene function was used to identify the variable genes for dimensionality reduction using PCA function. UMAP was used for dimensionality reduction and clustering of the cells.

Revised version: We used the standard Seurat v4.2 analyses pipeline to analyze our data (89). In brief the low-quality cells and potential doublets were excluded from the dataset to create the analyses dataset. The residual features due to the presence of RBCs were corrected before integration. “scTransform” function was used for integration of all the datasets after which the standard Seurat pipeline was used for data normalization of integrated data. “FindVariableGenes” gene function was used to identify the variable genes for dimensionality reduction using PCA function. UMAP was used for dimensionality reduction and clustering of cells.

Earlier version: To identify the unique features and cell clusters within each cell subtypes, we used the sub-setting feature within Seurat. After identifying the 5 major cell populations (epithelial, endothelial, stromal, myeloid and lymphoid) in our data sets, each of these cell types were sub-clustered using “subset ()” function, normalized and re-clustered. Cell annotation for each of the subsets was performed with the help of Tabula Muris database (35, 96). However, some clusters were annotated manually with the help of a literature search as discussed in the results section.

Revised version: To identify the unique features and cell clusters within each cell subtype, we used the sub-setting feature within Seurat. After identifying the 5 major cell populations (epithelial, endothelial, stromal, myeloid and lymphoid) in our data sets, each of these cell types were sub-clustered using “subset” function, normalized and re-clustered. Cell annotation for each of the subsets was performed with the help of Tabula Muris database (90, 91). However, some clusters were annotated manually with the help of a literature search as discussed in the results section.

Earlier version: DESeq2 (97) was used to perform pseudobulk analyses to identify differentially expressed genes within each group. The ClusterProfiler (98) was employed to perform gene enrichment analyses of the differentially expressed genes.

Revised version: DESeq2 (V. 1.42.1) was used to perform pseudobulk analyses to identify differentially expressed genes within each group. The ClusterProfiler R package (V. 4.10.1) (92) was employed to perform gene enrichment analyses of the differentially expressed genes.

Cytokine/chemokine assessment

Earlier version: We used multiplex assay to determine the levels of cytokine/chemokine in the lung homogenates from control and e-cig aerosol exposed mouse lungs using commercially available Bio-Plex Pro Mouse Chemokine Assay (Cat#12009159, Bio-RAD, Hercules, CA) per the manufacturer’s instructions. Approximately 40 mg of mouse lung lobes were homogenized in 300uL of 1X RIPA buffer with 0.1% protease and phosphatase inhibitor. The lung homogenate was stored on ice for 30 min. Following incubation, the homogenate was centrifuged at 15000 rpm for 15 min at 4 degrees Celsius. The supernatant was collected and used for performing the multianalyte assay for determination of cytokine/chemokine levels using Luminex FlexMap3D system

Revised version: We used multiplex assay to determine the levels of cytokine/chemokine in the lung homogenates from control and e-cig aerosol exposed mouse lungs using commercially available Bio-Plex Pro Mouse Chemokine Assay (Cat#12009159, Bio-RAD, Hercules, CA) per the manufacturer’s instructions. Approximately 40 mg of mouse lung lobes were homogenized in 300uL of 1X RIPA buffer with 0.1% protease and phosphatase inhibitor. The lung homogenate was stored on ice for 30 min. Following incubation, the homogenate was centrifuged at 15000 rpm for 15 min at 4 degrees Celsius. The supernatant was collected and used for performing the multianalyte assay for determination of cytokine/chemokine levels using Luminex FlexMap3D system. A heatmap after normalization of the measured cytokine/chemokine to the protein amount loaded was plotted.

Lung Histology

Earlier version: The left lung lobe of mice used for scRNA seq were inflated with 1% low melting agarose and fixed with 4% neutral buffered PFA. Fixed lungs were dehydrated, before being paraffin-embedded and sectioned (5μm). Hematoxylin and eosin (H&E) staining was performed by the Histology, Biochemistry, and Molecular Imaging Core at URMC. We employed MetaMorph software (Molecular Devices) to determine the “Mean linear intercept” (Lm) of airspace. In brief, 6-10 randomly selected ×100 fields per slide were photographed in a blinded manner, and Lm values were quantitated for each captured image using the MetaMorph software. Care was taken to avoid pictures with bronchiolar regions and knife marks in the tissue for quantitation. An average of all quantified images was then plotted for each sample in a sex-dependent manner.

Revised version: The left lung lobe of mice used for scRNA seq were inflated with 1% low melting agarose and fixed with 4% neutral buffered PFA. Fixed lungs were dehydrated, before being paraffin-embedded and sectioned (5μm). Hematoxylin and eosin (H&E) staining was performed by the Histology, Biochemistry, and Molecular Imaging Core at URMC. The H&E stain was observed at 10X magnification using Nikon Elipse-Ni fluorescence microscope. Ten to fifteen random images were captured per sample.

Flow cytometry

Earlier version: Flow cytometry was performed on the cells collected from BALF and lung homogenates from air and flavored e-cig aerosol exposed mouse lungs. For analyses of cells collected from BALF, the fluid was centrifuged at 2000g for 10 min at 40C. Supernatant was collected for ELISA, while the cell pellet was re-suspended in 1mL of PBS. For analyses of immune cell population in the lung, the lung lobes were digested as described earlier (94). The single-cell suspension thus prepared was used to run flow cytometry using the BD LSRFortessa cell analyzer. Cells were blocked with CD16/32 (Tonbo biosciences 70-0161-u500, 1:10) to prevent nonspecific binding and stained with a master mix of Siglec F (BD OptiBuild Cat#740280, 1:200), CD11b (Biolegend Cat #101243, 1:200), Ly6G (BD Horizon Cat# 562700, 1:200), CD45 (Biolegend Cat#103126, 1:200), CD11c (Biolegend Cat #117318, 1:200), CD4 (Biolegend Cat#116012, 1:200), and CD8 (eBiosciences Cat#17-0081-82, 1:200). 7AAD (eBiosciences Cat#00-6993-50, 1:10) was used as the nucleic acid dye to detect live and dead cells.

Revised version: Flow cytometry was performed on the cells collected from lung homogenates from air and flavored e-cig aerosol exposed mouse lungs. For analyses of immune cell population in the lung, the lung lobes were digested as described earlier (88). The single-cell suspension thus prepared was used to run flow cytometry using the BD LSRFortessa cell analyzer. Cells were blocked with CD16/32 (Tonbo biosciences 70- 0161-u500, 1:10) to prevent nonspecific binding and stained with a master mix of Siglec F (BD OptiBuild Cat#740280, 1:200), CD11b (Biolegend Cat #101243, 1:200), Ly6G (BD Horizon Cat# 562700, 1:200), CD45 (Biolegend Cat#103126, 1:200), CD11c (Biolegend Cat #117318, 1:200), CD4 (Biolegend Cat#116012, 1:200), and CD8 (eBiosciences Cat#17-0081-82, 1:200). 7AAD (eBiosciences Cat#00-6993-50, 1:10) was used as the nucleic acid dye to detect live and dead cells.

Metal analyses

To understand the levels of metals released during subsequent days of exposure, we performed Inductively coupled plasma mass spectrometry (ICP-MS) on the e-cig aerosol condensates collected from each day of exposure using a Perkin Elmer ICP-MS model 2000C. The samples were run using a Total Quant KED protocol with 4mL/min Helium flow and externally calibrated using a blank and a 100ppb standard for the 51 elements. The samples were submitted to the Element Analyses facility at URMC, and levels of metals thus detected were plotted.

Added in the revised version: Ly6G /S100A8 Double staining

Added in the revised version: To determine the presence of Ly6G+ and Ly6G- neutrophils, FFPE tissue sections from air and tobacco-flavored e-cig aerosol exposed lungs were stained with Ly6G and S100A8. In brief, 2-3 tissue sections per sample were deparaffinized using serial incubation in xylene followed by graded alcohol. Slides were incubated in 1X Citrate Buffer (Cat# S1699, Agilent, Santa Clara, CA) for 10 min at 95°C for antigen retrieval which was followed by incubation at room temperature for 30 min. The slides were next washed with water and permeabilized using a permeabilization buffer (0.1% Triton-X in 1X TBST) for 10 min. Next, the slides were again washed with 1X TBST and blocked using Blocking buffer (5% goat serum in 1X TBST) for 30 min at room temperature. The blocked slides were incubated overnight at 4°C with Ly6G (Cat# 16-9668-85, Invitrogen, dilution: 1:100) and S100A8 (Cat# 26992-1-AP, Proteintech, dilution: 1:200). Next day, the slides were washed with 1X TBST and incubated for 2 hrs. at room temperature with goat anti-rabbit Alexa Fluor 594 (Cat # A11012, Invitrogen) and donkey anti-mouse Alexa Fluor 488 (Cat # A21202, Invitrogen) secondary antibody at 1:1000 dilution. Thereafter the slides were washed and mounted with ProLong Diamond Antifade Mountant with DAPI (Cat# P36962, Invitrogen, Waltham, MA). 6-10 images were captured using Cytation 5 Imaging software (Thermo Fisher, Waltham, MA) at 10X magnification.

Statistical significance

Earlier version: We used GraphPad Prism 10.0 for all statistical calculations. All the data plotted in this paper are expressed as mean ± SEM. Pairwise comparisons were done using unpaired t test while one-way analysis of variance (ANOVA) with ad-hoc Tukey’s test was employed for multi-group comparisons. To identify sex-based variations in our treatment groups, Tukey post hoc two-way ANOVA was employed.

Revised version: We used GraphPad Prism 10.5.0 for all statistical calculations. All the data plotted in this paper are expressed as mean ± SEM. Pairwise comparisons were done using unpaired t test while one-way analysis of variance (ANOVA) with ad-hoc Tukey’s test was employed for multi-group comparisons. To identify sex-based variations in our treatment groups, Tukey post hoc two-way ANOVA was employed.